Estradiol elicits distinct firing patterns in arcuate nucleus kisspeptin neurons of females through altering ion channel conductances

Author response:

The following is the authors’ response to the current reviews.

Reviewer #1 (Public review):

Summary:

In this work, Qiu and colleagues examined the effects of preovulatory (i.e., proestrous or late follicular phase) levels of circulating estradiol on multiple calcium and potassium channel conductances in arcuate nucleus kisspeptin neurons. Although these cells are strongly linked to a role as the "GnRH pulse generator," the goal here was to examine the physiological properties of these cells in a hormonal milieu mimicking late proestrus, the time of the preovulatory GnRH-LH surge. Computational modeling is used to manipulate multiple conductances simultaneously and support a role for certain calcium channels in facilitating a switch in firing mode from tonic to bursting. CRISPR knockdown of the TRPC5 channel reduced overall excitability, but this was only examined in cells from ovariectomized mice without estradiol treatment.

Comments to address most recent author response:

The concern regarding the CRISPR experiments being confined to OVX mice is that the results can only suggest that CRISPR-mediated knockdown of TRPC5 can, at best, phenocopy the OVX+E condition. A reciprocal experiment in the opposite direction (for example, that returning TRPC5 to OVX levels in OVX+E mice prevents the changes in firing activity and pattern typical of the OVX+E2 condition) would strengthen the indication that E2-sensitive changes in TRPC5 expression and function are critically important to surge function. Acknowledging this as a limitation of the studies would help to better contextualize the value of the CRISPR experiments to an understanding of surge mechanisms when done only in OVX conditions.

We have noted in the manuscript that “It would be of interest in future experiments to do the reciprocal experiment to see if overexpressing Trpc5 channels in Kiss1ARH neurons from OVX + E2 females restores the RMP and “rescues” the synchronization phenotype.”

The nature of the confusion regarding the consideration of OVX+E2 conditions in the computational model primarily arises from the methods description in the supplemental file: "The effect of E2 on ionic currents is modelled as a change in the maximum conductance parameter. For currents IM,IT, ICa and ITRPC5 this change is inferred from the qPCR data assuming that the conductance is directly proportional to the mRNA expression." If these were instead based on the whole-cell recordings as the authors now indicate in their response, then this description needs to be edited and clarified accordingly. Furthermore, the section states, "For ISK, IBK, Ileak, the OVX and OVX+E2 conductances are obtained from current-voltage relationships recorded from Kiss1ARH neurons in the absence/presence of iberiotoxin (BK blocker) and apamin (SK blocker). All other currents were assumed to be unaffected by E2." This section thus does not directly indicate that the recordings in the stated figures were used in the model, and moreover suggests that currents besides ISK, IBK, and Ileak were not different in OVX+E2 conditions.

The prior evidence stated for correlation of mRNA and channel conductance is not explicitly cited in the manuscript. It is well known that post-translational modifications, physiological modulation of individual channel biophysical properties, and many other factors can influence the end output of a membrane conductance. Therefore, the authors should, at minimum, provide a literature citation supporting the assumption used here.

We have re-written the paragraph on “Modelling the effects of E2” in the Supplemental Information (now Appendix 1) to clarify the that the modeling was based on a combination of electrophysiological recordings and the qPCR data presented in this and previous publications. The statement that “all other currents were assumed to be unaffected by E2” was a misstatement and has been deleted. As per the reviewer’s request, we have listed seven publications that document the correlation between the mRNA expression and channel conductance for the various channels. We thank the reviewer for the suggestion.

Reviewer #2 (Public review):

Summary:

Kisspeptin neurons of the arcuate nucleus (ARC) are thought to be responsible for the pulsatile GnRH secretory pattern and to mediate feedback regulation of GnRH secretion by estradiol (E2). Evidence in the literature, including the work of the authors, indicates that ARC kisspeptin coordinate their activity through reciprocal synaptic interactions and the release of glutamate and of neuropeptide neurokinin B (NKB), which they co-express. The authors show here that E2 regulates the expression of genes encoding different voltage-dependent calcium channels, calcium-dependent potassium channels and canonical transient receptor potential (TRPC5) channels and of the corresponding ionic currents in ARC kisspeptin neurons. Using computer simulations of the electrical activity of ARC kisspeptin neurons, the authors also provide evidence of what these changes translate into in terms of these cells' firing patterns. The experiments reveal that E2 upregulates various voltage-gated calcium currents as well as 2 subtypes of calcium-dependent potassium currents, while decreasing TRPC5 expression (an ion channel downstream of NKB receptor activation), the slow excitatory synaptic potentials (slow EPSP) elicited in ARC kisspeptin neurons by NKB release and expression of the G protein-associated inward-rectifying potassium channel (GIRK). Based on these results, and on those of computer simulations, the authors propose that E2 promotes a functional transition of ARC kisspeptin neurons from neuropeptide-mediated sustained firing that supports coordinated activity for pulsatile GnRH secretion to a less intense burst-like firing pattern that could favor glutamate release from ARC kisspeptin. The authors suggest that the latter might be important for the generation of the preovulatory surge in females.

Strengths:

The authors combined multiple approaches in vitro and in silico to gain insights into the impact of E2 on the electrical activity of ARC kisspeptin neurons. These include patch-clamp electrophysiology combined with selective optogenetic stimulation of ARC kisspeptin neurons, reverse transcriptase quantitative PCR, pharmacology and CRISPR-Cas9-mediated knockdown of the Trpc5 gene. The addition of computer simulations for understanding the impact of E2 on the electrical activity of ARC kisspeptin cells is also a strength.

The authors add interesting information on the complement of ionic currents in ARC kisspeptin neurons and on their regulation by E2 to what was already known in the literature. Pharmacological and electrophysiological experiments appear of the highest standards and robust statistical analyses are provided throughout. The impact of E2 replacement on calcium and potassium currents is compelling. Likewise, the results of Trpc5 gene knockdown do provide good evidence that the TRPC5 channel plays a key role in mediating the NKB-mediated slow EPSP. Surprisingly, this also revealed an unsuspected role for this channel in regulating the membrane potential and excitability of ARC kisspeptin neurons.

Weaknesses:

The manuscript also has weaknesses that obscure some of the conclusions drawn by the authors.

One is that the authors compare here two conditions, OVX versus OVX replaced with high E2, that may not reflect the physiological conditions under which the proposed transition between neuropeptide-dependent sustained firing and less intense burst firing might take place (i.e. the diestrous [low E2] and proestrous [high E2] stages of the estrous cycle). This is an important caveat to keep in mind when interpreting the authors' findings. Indeed, that E2 alters certain ionic currents when added back to OVX females, does not mean that the magnitude of all of these ionic currents will vary during the estrous cycle.

We do know that the slow EPSP, which is generated by TRPC5 channels, tracks beautifully with the steroid state of female mice. Using our E2 treatment paradigm that generates a LH surge in OVX females (left panel in Author response image 1), there is no difference in the amplitude of the slow EPSP in proestrous versus OVX + E2 females (right panel in Author response image 1).

In addition, although the computational modeling indicates a role of the various E2-modulated conductances in causing a transition in ARC kisspeptin neuron firing pattern, their role is not directly tested in physiological recordings, weakening the link between these changes and the shift in firing patterns.

In future experiments we will test directly the physiological contribution of the other E2-modulated conductances in causing the transition in the firing pattern of arcuate Kiss1 neurons using CRISPR/SaCas9 technology as we have documented for the TRPC5 channel (e.g., Figures 11 and 12).

Overall, the manuscript provides interesting information about the effects of E2 on specific ionic currents in ARC kisspeptin neurons and some insights into the functional impact of these changes. However, some of the conclusions of the work, with regard, in particular, to the role of these changes in ion channels and to their implications for the LH surge, are not fully supported by the findings.

---------

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

In this work, Qiu and colleagues examined the effects of preovulatory (i.e., proestrous or late follicular phase) levels of circulating estradiol on multiple calcium and potassium channel conductances in arcuate nucleus kisspeptin neurons. Although these cells are strongly linked to a role as the "GnRH pulse generator," the goal here was to examine the physiological properties of these cells in a hormonal milieu mimicking late proestrus, the time of the preovulatory GnRH-LH surge. Computational modeling is used to manipulate multiple conductances simultaneously and support a role for certain calcium channels in facilitating a switch in firing mode from tonic to bursting. CRISPR knockdown of the TRPC5 channel reduced overall excitability, but this was only examined in cells from ovariectomized mice without estradiol treatment. The manuscript has been substantially improved from the initial version by the addition of new experiments and clarification of important figures. Importantly, the overlap of data with previous reports from the same group has been corrected.

Strengths:

(1) Examination of multiple types of calcium and potassium currents, both through electrophysiology and molecular biology.

(2) Focus on arcuate kisspeptin neurons during the surge is relatively conceptually novel as the anteroventral periventricular nucleus (AVPV) kisspeptin neurons have received much more attention as the "surge generator" population.

(3) The modeling studies allow for direct examination of manipulation of single and multiple conductances, whereas the electrophysiology studies necessarily require examination of each current in isolation. Construction of an arcuate kisspeptin neuron model promises to be of value to the reproductive neuroendocrinology field.

Weaknesses:

A remaining weakness in this revised version of the manuscript is that the relevance of the CRISPR experiments is still rather tenuous given that the goal is to understand what happens in the estrogen-treatment condition, and these experiments were performed only in OVX mice. Similar concerns reflect that the computational model examining the effect of E2 infers multiple conductances based on qPCR data and an assumption that the conductances are directionally proportional to the level of gene expression, and then tunes these to the current recordings obtained from OVX mice, without a direct confirmation in OVX+E2 conditions that the model parameters accurately reflect the properties of these currents in the presence of estrogen.

We are still puzzled by Reviewer’s concerns about doing the CRISPRing of Trpc5 in the OVX+E2 females. The Trpc5 channel expression is significantly reduced with the E2 treatment (Figure 10E) which we know translates into a minimal slow EPSP (Figure 2, Qiu eLife 2016) and is essentially equivalent to the slow EPSP amplitude in the Trpc5 mutagenesis in the ovariectomized females (Figure 12). TRPC5 channel conductance is already at “rock bottom.” The modeling informs us that such a low TRPC5 conductance will not support a long lasting slow EPSP and sustained firing (Figure 13A).

Also, we respectively point out that we have published a score of papers over the past 20 years showing that the channel conductance does correlate with the mRNA expression (e.g., Qiu et al., eLife 2018). Secondly, the model does take into consideration the OVX + E2 conditions (Figure 13B,C) which is based on the extensive whole-cell recordings presented in Figures 4,5,6,7,8 and 9.

Reviewer #2 (Public Review):

Summary:

Kisspeptin neurons of the arcuate nucleus (ARC) are thought to be responsible for the pulsatile GnRH secretory pattern and to mediate feedback regulation of GnRH secretion by estradiol (E2). Evidence in the literature, including the work of the authors, indicates that ARC kisspeptin coordinate their activity through reciprocal synaptic interactions and the release of glutamate and of neuropeptide neurokinin B (NKB), which they co-express. The authors show here that E2 regulates the expression of genes encoding different voltage-dependent calcium channels, calcium-dependent potassium channels and canonical transient receptor potential (TRPC5) channels and of the corresponding ionic currents in ARC kisspeptin neurons. Using computer simulations of the electrical activity of ARC kisspeptin neurons, the authors also provide evidence of what these changes translate into in terms of these cells' firing patterns. The experiments reveal that E2 upregulates various voltage-gated calcium currents as well as 2 subtypes of calcium-dependent potassium currents while decreasing TRPC5 expression (an ion channel downstream of NKB receptor activation), the slow excitatory synaptic potentials (slow EPSP) elicited in ARC kisspeptin neurons by NKB release and expression of the G protein-associated inward-rectifying potassium channel (GIRK). Based on these results, and on those of computer simulations, the authors propose that E2 promotes a functional transition of ARC kisspeptin neurons from neuropeptide-mediated sustained firing that supports coordinated activity for pulsatile GnRH secretion to a less intense burst-like firing pattern that could favor glutamate release from ARC kisspeptin. The authors suggest that the latter might be important for the generation of the preovulatory surge in females.

Strengths:

The authors combined multiple approaches in vitro and in silico to gain insights into the impact of E2 on the electrical activity of ARC kisspeptin neurons. These include patch-clamp electrophysiology combined with selective optogenetic stimulation of ARC kisspeptin neurons, reverse transcriptase quantitative PCR, pharmacology and CRISPR-Cas9-mediated knockdown of the Trpc5 gene. The addition of computer simulations for understanding the impact of E2 on the electrical activity of ARC kisspeptin cells is also a strength.

The authors add interesting information on the complement of ionic currents in ARC kisspeptin neurons and on their regulation by E2 to what was already known in the literature. Pharmacological and electrophysiological experiments appear of the highest standards and robust statistical analyses are provided throughout. The impact of E2 replacement on calcium and potassium currents is compelling. Likewise, the results of Trpc5 gene knockdown do provide good evidence that the TRPC5 channel plays a key role in mediating the NKB-mediated slow EPSP. Surprisingly, this also revealed an unsuspected role for this channel in regulating the membrane potential and excitability of ARC kisspeptin neurons.

Weaknesses:

The manuscript also has weaknesses that obscure some of the conclusions drawn by the authors.

One is that the authors compare here two conditions, OVX versus OVX replaced with high E2, that may not reflect the physiological conditions under which the proposed transition between neuropeptide-dependent sustained firing and less intense burst firing might take place (i.e. the diestrous [low E2] and proestrous [high E2] stages of the estrous cycle). This is an important caveat to keep in mind when interpreting the authors' findings. Indeed, that E2 alters certain ionic currents when added back to OVX females, does not mean that the magnitude of all of these ionic currents will vary during the estrous cycle.

Unfortunately, mice are a poor reproductive model since female mice do not have a clear follicular (estradiol-driven) phase distinctive from the luteal (progesterone-driven) phase. Had we utilized a “proestrous” female, we could not with certainty distinguish between the effects of estradiol versus progesterone on the expression of the calcium and potassium channels that were the focus of this study. Therefore, using our physiological model we can state with confidence that “estradiol elicits distinct firing patterns in arcuate nucleus kisspeptin neurons….”

Overall, the manuscript provides interesting information about the effects of E2 on specific ionic currents in ARC kisspeptin neurons and some insights into the functional impact of these changes. However, some of the conclusions of the work, with regard, in particular, to the role of these changes in ion channels and their implications for the LH surge, are not fully supported by the findings.

As we pointed out in the Discussion, the O’Byrne lab has clearly shown the relevance of Kiss1ARH neuronal burst firing and the release of glutamate to its effects on the LH surge:

“Rather, we postulate that glutamate neurotransmission is more important for excitation of Kiss1AVPV/PeN neurons and facilitating the GnRH (LH) surge with high circulating levels of E2 when peptide neurotransmitters are at a nadir and glutamate levels are high in female Kiss1ARH neurons. Indeed, low frequency (5 Hz) optogenetic stimulation of Kiss1ARH neurons, which only releases glutamate in E2-treated, ovariectomized females (Qiu J. et al., 2016), generates a surge-like increase in LH release during periods of optical stimulation (Lin et al., 2021; Voliotis et al., 2021). In a subsequent study optical stimulation of Kiss1ARH neuron terminals in the AVPV at 20 Hz, a frequency commonly used for terminal stimulation in vivo, generated a similar surge of LH (Shen et al., 2022). Additionally, intra-AVPV infusion of glutamate antagonists, AP5+CNQX, completely blocked the LH surge induced by Kiss1ARH terminal photostimulation in the AVPV (Shen et al., 2022).”

Recommendations for the authors:

Reviewer #2 (Recommendations for The Authors):

The reviewer noted the following in the revised manuscript:

- page 6, the authors may consider adding that presynaptic effects of blocking calcium channels on the slow EPSP cannot be fully ruled out. Indeed, the added experiments do indicate that some of the effects can be explained by impaired regulation of TRPC5 channels by calcium influx through calcium channels; however, the senktide-induced current is not fully blocked by the broad-spectrum calcium channel inhibitor cadmium, suggesting that the effect of blocking these channels on the slow EPSP may involve other mechanisms, such as presynaptic effects.

Optogenetic stimulation of all Kiss1ARH neurons induces the release of NKB at “physiological” concentrations, which in turn generates a slow EPSP in the recorded Kiss1ARH neuron. Blocking voltage-gated calcium channels can inhibit the NKB release from presynaptic Kiss1ARH neurons, thereby reducing the amplitude of the slow EPSP. However, in whole-cell recordings of synaptically isolated Kiss1ARH neurons, senktide directly induces a large inward current (Figure 3F), which is generated by the opening of TRPC5 channels (Qiu et al. J. Neurosci 2021). Voltage-gated calcium channels are coupled to the activation of TRPC5 channels (Blair, Kaczmarek and Clapham, J. Gen Physiol 2009), so by blocking voltage-gated calcium channels, cadmium effectively abrogates the facilitating effects of these channels on TRPC5 channel activation and significantly reduces but does not abolish the inward (excitatory) current (Figures 3F-H). We have clarified in the Results (page 6) that the Kiss1ARH neurons were synaptically isolated as depicted in Figures 3F,G.

- page 8, bottom, the mean value given for the apamin-sensitive current amplitude in E2 treated females does not match that plotted on the I/V graph in Figure 7F.

Thank you for pointing out this typographical error, which we have corrected.

Changes from the earlier version to the revised version

Abstract

Abstract

Hypothalamic kisspeptin (Kiss1) neurons are vital for pubertal development and reproduction. Arcuate nucleus Kiss1 (Kiss1ARH) neurons are responsible for the pulsatile release of Gonadotropin-releasing Hormone (GnRH). In females, the behavior of Kiss1ARH neurons, expressing Kiss1, Neurokinin B (NKB), and Dynorphin (Dyn), varies throughout the ovarian cycle. Studies indicate that 17β-estradiol (E2) reduces peptide expression but increases Vglut2 mRNA and glutamate neurotransmission in these neurons, suggesting a shift from peptidergic to glutamatergic signaling. To investigate this shift, we combined transcriptomics, electrophysiology, and mathematical modeling. Our results demonstrate that E2 treatment upregulates the mRNA expression of voltage-activated calcium channels, elevating the whole-cell calcium current and that contribute to high-frequency burst firing. Additionally, E2 treatment decreased the mRNA levels of Canonical Transient Receptor Potential (TPRC) 5 and G protein-coupled K+ (GIRK) channels. When TRPC5 channels in Kiss1ARH neurons were deleted using CRISPR, the slow excitatory postsynaptic potential (sEPSP) was eliminated. Our data enabled us to formulate a biophysically realistic mathematical model of the Kiss1ARH neuron, suggesting that E2 modifies ionic conductances in Kiss1ARH neurons, enabling the transition from high frequency synchronous firing through NKB-driven activation of TRPC5 channels to a short bursting mode facilitating glutamate release. In a low E2 milieu, synchronous firing of Kiss1ARH neurons drives pulsatile release of GnRH, while the transition to burst firing with high, preovulatory levels of E2 would facilitate the GnRH surge through its glutamatergic synaptic connection to preoptic Kiss1 neurons.

Introduction

Hypothalamic kisspeptin (Kiss1) neurons and its cognate receptor (GPR 54 or Kiss1 R) are essential for pubertal development and reproduction, and may also be involved in the control of energy homeostasis (Kotani et al., 2001) (De Roux et al., 2003) (Seminara et al., 2003) (Messager et al., 2005) (Shahab et al., 2005) (d’Anglemont de Tassigny et al., 2008) (Qiu J. et al., 2018) (Rønnekleiv et al., 2022). Kisspeptin neurons within the arcuate nucleus of the hypothalamus (Kiss1ARH) co-express Kiss1, Neurokinin B (NKB) and Dynorphin (Dyn), which are all down-regulated by 17β-estradiol (E2) (Goodman et al., 2007; Navarro V. M. et al., 2009). However, Kiss1ARH neurons also express vesicular glutamate transporter 2 (vGlut2) and release glutamate, and both vGlut2 expression and glutamate release are upregulated by E2 in females (Qiu J. et al., 2018). This indicates that peptides and glutamate in Kiss1ARH neurons are differently modulated by E2 and suggests that there is a complex E2 regulation in these neurons such that they transition from predominantly peptidergic to glutamatergic neurotransmission and hence from a “pulsatile” to “surge” mode. It has been known for decades that neurons located within the arcuate nucleus are responsible for the pulsatile release of GnRH and subsequently pulsatile release of LH from the pituitary gland (O’Byrne et al., 1991) (Moenter et al., 1993). In this respect, it is now generally accepted that Kiss1ARH neurons are the main neurons responsible for the generation of pulsatile LH release, but the underlying cellular conductances generating this activity have not been elucidated.

Earlier version: Morphological studies have provided evidence that Kiss1ARH neurons can communicate directly with each other (Lehman et al., 2010) (Navarro V. M. et al., 2009) (Navarro V.M. et al., 2011). Furthermore, Kiss1ARH neurons express the NKB receptor, Tacr3, as well as the kappa (κ) opioid receptor (KOR), whereas GPR54 is not expressed in Kiss1ARH neurons, rendering them unresponsive to kisspeptin (d’Anglemont de Tassigny et al., 2008; Navarro V. M. et al., 2009; Wakabayashi et al., 2010). Using optogenetics and whole-cell recordings we demonstrated that high-frequency photoactivation of Kiss1ARH neurons induces an NKB-mediated slow excitatory postsynaptic potential (EPSP), which is mediated by the recruitment of canonical transient receptor potential (TRCP5) channels. The release of NKB is limited by co-released dynorphin, which acts presynaptically to inhibit further release. Together the two peptides cause synchronized firing of Kiss1ARH neurons (Kelly et al., 2018; Qiu J. et al., 2016; Qiu Jian et al., 2021), whereas kisspeptin and glutamate appear to be the main output signals from Kiss1ARH neurons (Qiu J. et al., 2016; Qiu J. et al., 2018) (Voliotis et al., 2021) (Liu et al., 2021).

Revised version: Morphological studies have provided evidence that Kiss1ARH neurons can communicate directly with each other (Lehman et al., 2010) (Navarro V. M. et al., 2009) (Navarro V.M. et al., 2011). Furthermore, Kiss1ARH neurons express the NKB receptor, Tacr3, as well as the kappa (κ) opioid receptor (KOR), whereas GPR54 is not expressed in Kiss1ARH neurons, rendering them unresponsive to kisspeptin (d’Anglemont de Tassigny et al., 2008; Navarro V. M. et al., 2009; Wakabayashi et al., 2010). Using optogenetics and whole-cell recordings we demonstrated that high-frequency photoactivation of Kiss1ARH neurons induces an NKB-mediated slow excitatory postsynaptic potential (EPSP), which is mediated by the recruitment of canonical transient receptor potential (TRCP5) channels. The release of NKB is limited by co-released dynorphin, which acts presynaptically to inhibit further release. Together the two peptides cause synchronized firing of Kiss1ARH neurons (Kelly et al., 2018; Qiu J. et al., 2016; Qiu Jian et al., 2021), whereas kisspeptin and glutamate appear to be the main output signals from Kiss1ARH neurons (Qiu J. et al., 2016; 3XXX) (Voliotis et al., 2021) (Liu et al., 2021).

Although single action potential-generated calcium influx is sufficient to trigger the release of classical neurotransmitters such as glutamate, high frequency (10-20 Hz) is required for the release of neuropeptides such as kisspeptin, NKB and dynorphin (Qiu J. et al., 2016). Indeed, the slow EPSP, which underlies the synchronization, is similar to the “plateau potential” that has been described in hippocampal and cortical neurons (Arboit et al., 2020; Zhang Z. et al., 2011). Many neurons, including Kiss1ARH neurons, express the biophysical properties that allow them to continue to persistently fire even after a triggering synaptic event has subsided (Zylberberg and Strowbridge, 2017) (Qiu J. et al., 2016). Moreover, the intrinsic bi-stability of neurons that generates persistent firing activity has been linked to a calcium-activated, non-selective cation current (ICAN) (Zylberberg and Strowbridge, 2017), and TRPC channels, specifically TRPC5 channels, are thought to be responsible for the ICAN in cortical neurons (Zhang Z. et al., 2011). Therefore, we postulated that TacR3 activation via NKB drives influx of Ca+2 through TRPC5 channels leading to greater build-up of [Ca2+]i that facilitates the opening of more TRPC5 channels in a self-sustaining manner. Indeed, using the fast intracellular calcium chelator BAPTA, which has been shown to robustly inhibit TRPC5 channel activation in heterologous cells (Blair et al., 2009), we have been able to abolish the slow EPSP and persistent firing in Kiss1ARH neurons following optogenetic stimulation in female mice (Qiu Jian et al., 2021).

Although the expression of peptides in Kiss1ARH neurons are downregulated by high circulating levels (late follicular levels) of E2, the intrinsic excitability of and the glutamate release by Kiss1ARH neurons are increased by Cacna1g (Cav3.1, T-type calcium channel), Hcn1 and Hcn2 (Hyperpolarization-activated, Cyclic Nucleotide Gated channels) mRNA expression and Vglut2 mRNA expression, respectively (Qiu J. et al., 2018). Burst firing in CNS neurons, which efficiently releases fast amino acid transmitters like glutamate, is generated primarily by the T-type calcium channel current (IT) (e.g., in thalamic relay neurons), and the rhythmicity of this burst firing is dependent on the h-current (Ih) (for review, see (Lüthi and McCormick, 1998; Zagotta and Siegelbaum, 1996)). Ih is mediated by the HCN channel family, which includes channel subtypes 1-4, of which Hcn1 and Hcn2 are the main channels in Kiss1ARH neurons. Ih depolarizes neurons from hyperpolarized states, raising the membrane potential into the range of IT activation (Erickson et al., 1993a, 1993b; Kelly and Rønnekleiv, 1994; Lüthi and McCormick, 1998; Zhang C. et al., 2009). IT is mediated by the low-threshold voltage-gated calcium channels, CaV3.1-3.3 (for review see (Perez-Reyes, 2003)). IT initiates a transient Ca2+-driven depolarization above the threshold for action potential initiation (i.e., a low threshold spike) (Llinás, 1988; Tsien et al., 1987). This depolarization then drives neurons to fire an ensemble (burst) of Na+-driven action potentials.

Based on the above compelling evidence we postulated that Kiss1ARH neurons transition from peptidergic neurotransmission, driving the pulsatile release of GnRH via kisspeptin release into the median eminence, to glutamatergic transmission that facilitates in the preovulatory surge of GnRH (Lin et al., 2021) (Shen et al., 2022) through their projection to the anteroventral periventricular/periventricular nucleus Kiss1 (Kiss1AVPV/PeN) neurons (Qiu J. et al., 2016). Therefore, we initiated studies to thoroughly characterize effects of high circulating (late follicular) levels of E2 on the expression of the full complement of voltage-activated calcium channels (and currents) and the opposing K+ channels, involved in the repolarization, on the excitability of Kiss1ARH neurons. We performed whole-cell recordings and single cell RT-PCR analysis of Kiss1ARH neurons to determine which channels are involved in the physiological transition from peptidergic to glutamatergic neurotransmission. Our physiological findings were incorporated into a mathematical model that accounts for the E2 effects on the firing activity of Kiss1ARH neurons and validates our hypothesis that high levels of E2 facilitate the transition from peptidergic to glutamatergic neurotransmission.

Results

Voltage-activated Ca2+ channels contribute to burst firing of Kiss1ARH neurons

Our whole-cell current clamp recordings of Kiss1ARH neurons revealed that there is a significant transition of these Kiss1 neurons from a silent/tonic firing mode in the ovariectomized state to bursting/irregular firing with E2 treatment (Figure 1). Therefore, we sought to elucidate the cationic channels contributing to this physiological transition. Previously, we have shown that an increase in the intracellular calcium concentrations can potentiate TRPC5 channel current in POMC neurons (Qiu J. et al., 2010). Additionally, chelating intracellular calcium with BAPTA abolishes the slow excitatory postsynaptic potential (EPSP) and persistent firing in Kiss1ARH neurons (Qiu Jian et al., 2021). Hence, to investigate the contributions of voltage-activated calcium channels (VGCCs) to the increase in intracellular calcium, we measured the peak calcium current contributed by both the low and high voltage-activated calcium channels. The inward currents evoked by the voltage pulses (150-ms pulses starting from a holding potential of -80 mV with 10-mV increments) were identified as calcium (Ca2+) currents, as they were blocked by the universal VGCC inhibitor Cd2+ (200 μM) (McNally et al., 2020) (Figures 2A-E). The maximum total inward and Cd2+-sensitive currents reached their peak amplitudes at -10 mV. To differentiate between various calcium channel subtypes present in Kiss1ARH neurons, we applied selective antagonists individually, allowing us to isolate the drug-sensitive current for each cell. When we individually applied specific antagonists, we observed partial inhibition of whole-cell Ca2+ currents. Treatment with 10 μM nifedipine (Hiraizumi et al., 2008; Kato et al., 2003; Lee et al., 2002), an L-type Ca2+ channel inhibitor, resulted in a partial inhibition (26.1%) of the whole-cell calcium current. Additionally, application of 2 μM ω-conotoxin GVIA (conoGVIA, targeting N-type channels) (Lee et al., 2002), 200 nM ω-agatoxin IVA (AgaIVA, targeting P/Q-type channels) (Kato et al., 2003; McNally et al., 2020), 100 nM SNX-482 (targeting R-type channels) (Hiraizumi et al., 2008), or 1 μM TTA-P2 (TTAP2, targeting T-type channels) (McNally et al., 2020) also led to partial inhibition of the Ca2+ current (Figure 2F). The observed reduction in current with each inhibitor indicates the presence of all the major subtypes of Ca2+ currents in Kiss1ARH neurons. Among the inhibitors used, the largest components of the whole-cell calcium current were sensitive to nifedipine (26.1%), conoGVIA (25.1%), and SNX-482 (31.1%) (Figures 2A, B, D and 2F). Subsequently, we documented contributions from TTA-P2-sensitive channels, accounting for approximately 6.7% of the total Ca2+ current, and AgaIVA-sensitive channels, which constituted approximately 3.9% (Figures 2E, C). These findings indicate that high voltage-activated L-, N-, and R-type channels constitute the largest component of the voltage-activated Ca2+ current in Kiss1ARH neurons that not only increases the overall excitability but also greatly facilitate TRPC5 channel opening (Blair et al., 2009), which is the major downstream target of TACR3 activation by NKB (Qiu J. et al., 2016).

Properties of the firing of Kiss1ARH neurons from OVX and E2-treated OVX mice.

Earlier version: A–E. Representative whole cell, current clamp recordings of spontaneous phasic burst firing (A, which is only seen in E2-treated mice with clear up and down states), irregular burst firing (B), tonic firing (C), irregular firing (D) and silent (E) in Kiss1ARH neurons from E2-treated OVX mice. F-G. Summary pie chart illustrating the distribution of firing patterns in Kiss1ARH neurons from OVX+E2 (F, from 6 mice) or OVX (G, from 8 mice) mice (OVX versus OVX+E2, Chi-square test, X2 = 16.05, df = 3, **p = 0.0011). H. Current clamp recording in a Kiss1ARH neuron from an OVX+E2 female demonstrating the response to glutamate (100 μM). RMP = -70 mV. Similar responses were observed in all recorded female Kiss1ARH neurons (n = 7 from 4 mice). I. The spiking activity above the bar in G was expanded to highlight the pronounced effects of glutamate on burst firing activity of Kiss1ARH neurons, characterized by an ensemble of spikes riding on top of low-threshold spikes (arrows). Drugs were rapidly perfused into the bath as a 4-μl bolus. J. Spontaneous AP frequency of Kiss1ARH neurons before, during, and after glutamate application (n = 7 cells). Data are presented as mean ± SEM. Statistical comparisons between different treatments were performed using one-way ANOVA (F(2, 18) = 14.60, p = 0.0002) followed by Bonferroni’s post hoc test. *** indicates p < 0.005. K-L. Representative traces showing the amplitude of slow EPSPs induced by high frequency photostimulation in Kiss1ARH neurons in the presence of ionotropic glutamate receptor antagonists CNQX and AP5 from OVX (K) and OVX+E2 (L) females. The arrows indicate the measurements of slow EPSP amplitude after low-pass filtering (shown in K). M. Bar graphs summarizing the slow EPSPs in the Kiss1ARH neurons from OVX and E2-treated OVX mice in the presence of CNQX and AP5. Statistical comparisons between the two groups were performed using an unpaired t-test (t(12) = 5.181, p =0.0002). Data are expressed as mean ± SEM , with data points representing the number of cells.

Revised version: A–E. Representative whole cell, current clamp recordings of spontaneous phasic burst firing (A, which is only seen in E2-treated mice with clear up and down states), irregular burst firing (B), tonic firing (C), irregular firing (D) and silent (E) in Kiss1ARH neurons from E2-treated OVX mice. F-G. Summary pie chart illustrating the distribution of firing patterns in Kiss1ARH neurons from OVX+E2 (F, from 6 mice) or OVX (G, from 8 mice) mice (OVX versus OVX+E2, X2(3)= 6.05, p = 0.0011). H. Current clamp recording in a Kiss1ARH neuron from an OVX+E2 female demonstrating the response to glutamate (100 μM). RMP = -70 mV. Similar responses were observed in all recorded female Kiss1ARH neurons (n = 7 from 4 mice). I. The spiking activity above the bar in G was expanded to highlight the pronounced effects of glutamate on burst firing activity of Kiss1ARH neurons, characterized by an ensemble of spikes riding on top of low-threshold spikes (arrows). Drugs were rapidly perfused into the bath as a 4-μl bolus. J. Spontaneous AP frequency of Kiss1ARH neurons before, during, and after glutamate application (n = 7 cells). Data are presented as mean ± SEM. Statistical comparisons between different treatments were performed using one-way ANOVA (F(2, 18) = 14.60, p = 0.0002) followed by Bonferroni’s post hoc test. *** indicates p < 0.005. K-L. Representative traces showing the amplitude of slow EPSPs induced by high frequency photostimulation in Kiss1ARH neurons in the presence of ionotropic glutamate receptor antagonists CNQX and AP5 from OVX (K) and OVX+E2 (L) females. The arrows indicate the measurements of slow EPSP amplitude after low-pass filtering (shown in K). M. Bar graphs summarizing the slow EPSPs in the Kiss1ARH neurons from OVX and E2-treated OVX mice in the presence of CNQX and AP5. Statistical comparisons between the two groups were performed using an unpaired t-test (t(12) = 5.181, p =0.0002). Data are expressed as mean ± SEM, with data points representing the number of cells.

Relative contribution of voltage-gated calcium currents in Kiss1ARH neurons from OVX mice.

A-E, representative current –voltage relationships showing that Cd2+ (non-selective blocker of calcium channels) sensitive peak currents were inhibited by different calcium channel blockers: A, nifedipine; B, ω-conotoxin GIVA; C, ω-agatoxin IVA; D, SNX-482; E, TTA-P2. F, the maximum peak currents were measured at -10 mV. The proportions of Ca2+ currents inhibited by nifidipine (L type), ω-conotoxin GVIA (N type), ω-agatoxin IVA (P/Q), SNX-482 (R type) and TTA-P2 (T type). Data are expressed as mean ± SEM, with data points representing the number of cells.

Voltage-activated Ca2+ channels contribute to generation of slow-EPSP in Kiss1ARH neurons

Earlier version: Our previous study utilizing optogenetics demonstrated that high frequency photostimulation of Kiss1ARH neurons releases NKB. This release of NKB induces slow excitatory postsynaptic potentials (EPSPs) that facilitate the recruitment of other Kiss1ARH neurons, resulting in synchronous firing of the Kiss1ARH neuronal population (Qiu J. et al., 2016). Additionally, chelating intracellular calcium with the fast chelator BAPTA abolishes the slow EPSP and persistent firing in Kiss1ARH neurons, highlighting the role of calcium signaling in these processes (Qiu Jian et al., 2021). To access the involvement of HVA channels in the generation of the slow EPSP, we conducted experiments where we blocked the L-type Ca2+ channels with nifedipine (10 μM) and the N- and P/Q-type Ca2+ channels with ω-conotoxin MVIIC (1 μM) (Nunemaker et al., 2003). We then measured the slow EPSP. Indeed, both nifedipine and ω-conotoxin MVIIC significantly inhibited the slow EPSP by 35.2% and 60.4%, respectively (Figure 3A, B, E). In addition, we used SNX (100 nM), which selectively blocks R-type Ca2+ channels, and the slow EPSP was reduced to 28.7% of its control value (Figure 3C, E). The selective T-channel blocker TTA-P2 (5 µM) inhibited the slow EPSP by 28.6% (Figure 3D, E). Since the pharmacological blockade of the various calcium channels could also affect the presynaptic release of neurotransmitters (i.e., NKB), we measured the direct response of Kiss1ARH neurons to the TACR3 agonist senktide in the presence or absence of cadmium, a universal calcium channel blocker (Figures 3F-H). Indeed, Cd2+ significantly inhibited the inward current induced by senktide, which supports our posit that the plasma membrane calcium channels contribute to the TRPC5 channels activation. Therefore, it appears that calcium channels contribute to maintaining the sustained depolarization underlying the slow EPSP.

Revised version: Our previous study utilizing optogenetics demonstrated that high frequency photostimulation of Kiss1ARH neurons releases NKB. This release of NKB induces slow excitatory postsynaptic potentials (EPSPs) that facilitates the recruitment of other Kiss1ARH neurons, resulting in synchronous firing of the Kiss1ARH neuronal population (Qiu J. et al., 2016). Additionally, chelating intracellular calcium with the fast chelator BAPTA abolishes the slow EPSP and persistent firing in Kiss1ARH neurons, highlighting the role of calcium signaling in these processes (Qiu Jian et al., 2021). To access the involvement of HVA channels in the generation of the slow EPSP, we conducted experiments where we blocked the L-type Ca2+ channels with nifedipine (10 μM) and the N- and P/Q-type Ca2+ channels with ω-conotoxin MVIIC (1 μM) (Nunemaker et al., 2003). We then measured the slow EPSP. Indeed, both nifedipine and ω-conotoxin MVIIC significantly inhibited the slow EPSP by 35.2% and 60.4%, respectively (Figure 3A, B, E). In addition, we used SNX (100 nM), which selectively blocks R-type Ca2+ channels, and the slow EPSP was reduced to 28.7% of its control value (Figure 3C, E). The selective T-channel blocker TTA-P2 (5 µM) inhibited the slow EPSP by 28.6% (Figure 3D, E). Since the pharmacological blockade of the various calcium channels could also affect the presynaptic release of neurotransmitters (i.e., NKB), we measured the direct response of Kiss1ARH neurons to the TACR3 agonist senktide in the presence or absence of cadmium, a universal calcium channel blocker. In synaptically isolated whole-cell recordings of Kiss1ARH neurons, Cd2+ significantly inhibited the inward current induced by senktide (Figures 3F-H), which supports our posit that plasma membrane voltage-activated calcium channels contribute to the TRPC5 channel activation. Therefore, it appears that calcium channels contribute to maintaining the sustained depolarization underlying the slow EPSP.

Blockade of HVA Ca2+ channels decreases the slow EPSP in Kiss1ARH neurons.

A-D, representative traces showing that the slow EPSPs induced by high frequency photostimulation were abolished by perfusing the blocker of the L-type calcium channel, nifedipine (A) or N- and P/Q -type calcium channels, ω-conotoxin MVIIC (B), or the R-type calcium channel, SNX 482 (C), or the T-type calcium channel, TTA-P2 (D), respectively. The arrows indicate the measurements of slow EPSP amplitude, denoted as R1 and R2, after low-pass filtering (shown in C). E. Bar graphs summarizing the effects of drugs on the R2/R1 ratios. The slow EPSP was generated in OVX Kiss1-Cre::Ai32 mice. Comparisons between different treatments were performed using a one-way ANOVA analysis (F (3, 51) = 14.36, p<0.0001) with the Bonferroni’s post hoc test. *, **, **** indicates p < 0.05, 0.01, 0.001, respectively vs. control. F-G. Representative traces show that in voltage clamp, senktide induced a significant inward current in Kiss1ARH neurons (F) in the presence of TTX. This current was inhibited by the calcium channel blocker CdCl2 (200 μM)(G) in another cell. H. Bar graphs summarize the effect of the calcium channel blocker CdCl2 on the senktide-induced inward currents. Un-paired t-test for F versus G: t(20) = 2.575, p = 0.0181; * p < 0.05. Data are expressed as mean ± SEM, with data points representing the number of cells.

E2 increases the mRNA expression and the whole-cell current of voltage-activated Ca2+ channels

The neuropeptides NKB (tachykinin2, Tac2) and kisspeptin (Kiss1), which are expressed in Kiss1ARH neurons, are crucial for the pulsatile release of gonadotropin-releasing hormone (GnRH) and reproductive processes. E2 decreases the expression of Kiss1 and Tac2 mRNA in Kiss1ARH neurons but enhances the excitability of Kiss1ARH neurons by amplifying the expression of Cacna1g, Hcn1, and Hcn2 mRNA, as well as increasing T-type calcium currents and h-currents, respectively (Qiu, 2018, 20613}. Moreover, E2 drives Slc17a6 mRNA expression and enhances glutamatergic synaptic input to arcuate neurons and Kiss1AVPV/PeN neurons (Qiu J. et al., 2018). As a result, the E2-driven increase in Kiss1ARH neuronal excitability and glutamate neurotransmission may play a crucial role in triggering the surge of GnRH, ultimately leading to the LH surge (Lin et al., 2021) (Shen et al., 2022).

To assess the impact of E2 on the modulation of voltage-activated calcium channels and Kiss1ARH neuronal excitability, we employed real-time PCR (qPCR) to measure the relative expression levels of ion channel subtypes in Kiss1ARH neurons. We compared the expression in E2-treated females to those treated with oil, using the specific primers listed in Table 1. The quantification was conducted on pools of 10 neurons, as indicated in the Methods and Materials. In both oil- and E2-treated females, we quantified the expression of Cav1.2 (L-type), Cav2.1 (P/Q-type), Cav2.2 (N-type) and Cav2.3 (R-type) mRNAs. Remarkably, all of these mRNA transcripts exhibited increased expression levels in response to E2 treatment (Figure 4). Congruent with our previous findings (Qiu J. et al., 2018), we observed that the mRNA expression of Cav3.1 and Hcn1 was also upregulated in response to E2 treatment (Figure 4). These results suggested that E2 has a regulatory effect on the expression and function of all of these ion channels in Kiss1ARH neurons. To determine whether the increased mRNA expression translated into functional changes at the cellular level, we measured the whole-cell calcium current in Kiss1ARH neurons obtained from ovariectomized mice treated with either vehicle or E2. We discovered that E2 treatment led to a significant increase in the peak calcium current density in Kiss1ARH neurons (Figure 5A-D), which was recapitulated by our computational modeling (Figure 5E). These findings indicate that the upregulation of the mRNA expression of calcium channels by E2 translated to an augmented peak calcium current in Kiss1ARH neurons.

A, E2 increases the expression of low and high voltage-activated calcium channels in Kiss1ARH neurons.

Kiss1ARH neurons (three 10-cell pools) were harvested from each of 5 vehicle- and 5 E2-treated, OVX females to quantify ion channel mRNA expression of low and high voltage activated calcium channels as described in the Methods. The analysis included: T-type (Cav3.1) low voltage-activated, as well as the following high-voltage activated channels: R-type (Cav 2.3), L-type (Cav 1.2), N-type (Cav 2.2) and P/Q-type (Cav 2.1) calcium channels. Interestingly, all of these channels were upregulated with E2 treatment, which significantly increased the whole-cell calcium current (see Figure 5). Bar graphs represent the mean ± SEM with data points representing individual animals (oil versus E2, Unpaired t test for Cav 3.1, t(8)= 7.105, ***p=0.0001; for Cav 2.3, t(8) = 3.007, *p = 0.0169; for Cav 1.2, t(8)= 2.721, *p =0.0262; for Cav 2.2, t(8) = 4.001, **p 0.0039; for Cav 2.1, t(8) = 4.225, **p =0.0028).B, The same Kiss1ARH neuronal pools were also analyzed for mRNA expression of HCN1 ion channels, and E2 also increased the expression of hyperpolarization-activated, cyclic-nucleotide gated HCN1 channels in Kiss1ARH neurons. HCN1 channel mRNA expression was the most highly upregulated by E2 treatment in Kiss1ARH neurons. The expression values were calculated via the ΔΔCT method, normalized to GAPDH and relative to the oil control values. (oil versus E2, Unpaired t test, t(8)= 11.450, ****p < 0.0001).

E2 treatment (positive-feedback regimen) increases the Ca2+ currents in Kiss1ARH neurons.

Earlier version: A-B, Ca2+ currents in Kiss1ARH neurons with the same membrane capacitance from oil-treated (A) or E2-treated (B) animals. C, the maximum peak currents were measured at -10 mV. The current amplitudes were normalized to the cell capacitance in all cases to calculate current density. The bar graphs summarized the density of Ca2+ current in Kiss1ARHneurons from oil-treated and E2-treated animals. The mean density was significantly greater in E2-treated (13.4 ± 0.9 pA/pF, n = 11) than in oil-treated OVX females (7.2 ± 0.5 pA/pF, n = 40) (unpaired t-test, t(49) = 5.75, ****p < 0.0001). D. Relative contribution of voltage-gated calcium currents in Kiss1ARH neurons from OVX, E2-treated mice. The maximum peak currents were measured at -10 mV. The proportions of Ca2+ currents inhibited by nifidipine (L type), ω-Conotoxin GVIA (N type), ω-agatoxin IVA (P/Q), SNX-482 (R type) and TTA-P2 (T type). Data are expressed as mean ± SEM, with data points representing the number of cells. E. The modeling predicts that E2-treated, OVX females exhibit a significantly greater inward Ca2+ current (red trace) than the vehicle-treated females (black trace). The coductance of the modelled voltage-gated calcium current (L-, N-, P/Q-, and R-type) is set to 2.1 nS in the OVX state and 2.8 nS in OVX+E2 state, while for the T-type is set to 0.66 nS in the OVX state and 5 nS in OVX+E2 state.

Revised version: A-B, Ca2+ currents in Kiss1ARH neurons with the same membrane capacitance from oil-treated (A) or E2-treated (B) animals. C, the maximum peak currents were measured at -10 mV. The current amplitudes were normalized to the cell capacitance in all cases to calculate current density. The bar graphs summarized the density of Ca2+ current in Kiss1ARH neurons from oil-treated and E2-treated animals. The mean density was significantly greater in E2-treated (13.4 ± 0.9 pA/pF, n = 11) than in oil-treated OVX females (7.2 ± 0.5 pA/pF, n = 40) (unpaired t-test, t(49) = 5.75, ****p < 0.0001). D. Relative contribution of voltage-gated calcium currents in Kiss1ARH neurons from OVX, E2-treated mice. The maximum peak currents were measured at -10 mV. The proportions of Ca2+ currents inhibited by nifidipine (L type), ω-Conotoxin GVIA (N type), ω-agatoxin IVA (P/Q), SNX-482 (R type) and TTA-P2 (T type). Data are expressed as mean ± SEM, with data points representing the number of cells. E. The modeling predicts that E2-treated, OVX females exhibit a significantly greater inward Ca2+ current (red trace) than the vehicle-treated females (black trace). The coductance of the modelled voltage-gated calcium current (L-, N-, P/Q-, and R-type) is set to 2.1 nS in the OVX state and 2.8 nS in OVX+E2 state, while for the T-type is set to 0.66 nS in the OVX state and 5 nS in OVX+E2 state.

The largest components of the calcium currents in Kiss1ARH neurons from the E2-treated, ovariectomized females were found to be sensitive to nifedipine, conoGVIA, and SNX-482, accounting for approximately 24.9%, 24.6%, and 27.0% of the total current across cells, respectively, which is very similar to their contributions to the whole-cell current from vehicle-treated, ovariectomized females (Figure 3). In addition, contributions from TTA-P2-sensitive channels accounted for approximately 11.1% of the total Ca2+ current, while agaIVA-sensitive channels contributed to approximately 11.0% (Figure 5D). These results highlight the prevalence of L-, N-, and R-type calcium channels as the major contributors to the whole-cell calcium current in Kiss1ARH neurons from E2-treated, ovariectomized females, but also the involvement of T-type and P/Q-type channels.

E2 does not alter the kinetics of calcium channel activation or de-inactivation

In order to determine whether the increase in peak voltage-activated calcium current density was the result of E2 regulating calcium channel kinetics, mRNA expression or both, we examined the voltage dependence of activation and inactivation. By measuring the voltage dependence of activation, we assessed how E2 affects the ability of calcium channels to open in response to membrane potential changes. Similarly, by examining the voltage dependence of inactivation, we determined how E2 influences the inactivation kinetics of calcium channels. Based on our results, there was no difference in the voltage dependence of activation between cells from the vehicle-treated control group (V1/2 = -32.3 ± 2.1 mV; n = 13) and cells from estrogen-treated females (V1/2 = -33.6 ± 2.5 mV; n = 11). Similarly, there was not a significant difference in the voltage dependence of inactivation between control cells (V1/2 = -48.9 ± 4.8 mV; n = 6) and estrogen-treated cells (V1/2 = -44.1 ± 1.9 mV; n = 5) (Figure 6). Therefore, although long-term E2 treatment increased the mRNA expression of HVA calcium channels (Figure 4), it did not affect the channel kinetics in Kiss1ARH neurons (Figure 6). Also, E2 downregulated the expression of Kcnd2 mRNA encoding Kv4.2, which is expressed in Kiss1ARH neurons (Mendonça et al., 2018) and has similar kinetics of activation as the T-type calcium channels (Oil-treated, ovariectomized females relative mRNA expression: 1.053, n = 5 animals versus E2-treated expression: 0.5643, n=5; t-test p = 0.0061). This opposing K+ current would dampen the inward calcium current. Therefore, it appears that long-term E2 treatment does not modulate calcium channel kinetics but rather alters the mRNA expression to increase the calcium channel conductance.

Earlier version: Voltage dependence of ICa in Kiss1ARHneurons from OVX and OVX+E2 mice.

Revised version: Voltage dependence of ICa in Kiss1ARH neurons from OVX and OVX+E2 mice.

A-B, top panels: activation and inactivation protocol. Bottom: representative traces. C-D, the mean V1/2 values for calcium channel activation were not significantly different for cells from controls versus cells from E2-treated females. Similarly, the V1/2 values for channel steady-state inactivation were similar for both groups. Data are expressed as mean ± SEM.

BK, SK and KCNQ channels are involved in modulating excitability of Kiss1ARH neurons

In the brain, calcium plays a crucial role in sculpting neuronal firing by activating potassium channels, which subsequently influence neuronal behavior (Nicoll, 1988; Storm, 1990). Since HVA and LVA calcium channels were expressed in Kiss1ARH neurons, all of which contribute to the elevation of intracellular calcium concentration ([Ca2+]i) that facilitates TRPC5 channel opening (Blair et al., 2009), our next step involved measuring the changes in Ca2+-activated K+ channel conductances and assessing their mRNA expression. In various cell types increases in cytosolic calcium levels, whether resulting from extracellular influx or intracellular release, lead to the activation of plasma membrane calcium-dependent potassium channels (Sah Pankaj and Louise Faber, 2002). Similarly, in Kiss1ARH neurons, these channels would be activated by calcium influx through all four types of high voltage-gated calcium channels, as well as the low voltage-activated calcium channel, which are all active during action potential firing. The activity of Ca2+-activated K+ channels play a crucial role in numerous physiological processes, including secretion and the regulation of neuronal firing properties. Two main families of Ca2+-activated K+ channel channels have been characterized, distinguished by their biophysical and pharmacological properties. These families are known as BK (Big Conductance K+) and SK (Small Conductance K+) channels in the CNS (Kshatri et al., 2018). BK channels are known for their high potassium selectivity and large single channel conductance, typically ranging from 100 to 300 pS. Activation of BK channels requires both calcium binding and membrane depolarization (Blatz and Magleby, 1987; Marty, 1989; Sah P., 1996; Storm, 1990). On the other hand, SK channels are simply activated by increases in cytosolic calcium levels, with their half-maximal activation at 0.3 μM (Bond et al., 1999).

Earlier version: To investigate K+ currents, the cells were maintained at a holding potential of -70 mV while being exposed to blockers CNQX, AP5, picrotoxin and TTX. Subsequently, the membrane potential was stepped by depolarizing voltages, ranging from -60 mV to +40 mV in 10 mV increments, for a duration of 500 ms (Brereton et al., 2013). This protocol was employed to activate K+ currents (Figure 7A). First, we examined SK currents. The mean current density was determined at the end of the voltage pulses. In vehicle-treated, OVX females the application of the SK channel blocker apamin (100 nM) (Spergel, 2007) led to a significant reduction in whole-cell currents in the +10 to +40 mV range (Figures 7A, B). The mean outward current density at +40 mV in the control group was 125.5 ± 7.2 pA/pF (n = 5) with the apamin-sensitive component contributing 60.2 ± 3.3 pA/pF (n = 5) (Figures 7A, C). In contrast, in the E2-treated females, the overall mean outward current density at +40 mV was 195.4 ± 14.6 pA/pF (n = 6), which was significantly greater than the vehicle control group (Figures 7D, E). However, there was no significant difference in the apamin-sensitive component between the vehicle-treated and E2-treated females, 60.2 ± 3.3 pA/pF (n = 5) versus 128.1 ± 13 pA/pF (n = 6), respectively (Figure 7F). Furthermore, to investigate the expression of the mRNAs encoding SK channel subunits in Kiss1ARH neurons from vehicle-treated and E2-treated OVX females, qPCR experiments were performed on 10-cell Kiss1ARH neuronal pools (Figure 7G). We focused on SK3 channels because these channels exhibit the highest expression in the hypothalamus and E2 regulates their expression (Bosch et al., 2002). E2 treatment had no effect on the mRNA expression of the SK3 subunit. These molecular findings support our electrophysiology results. Our computational model was calibrated so that SK channels contributed ∼50 pA/pF to the whole-cell outward K+ current in E2-treated females (Figure 7H).

Revised version: To investigate K+ currents, the cells were maintained at a holding potential of -70 mV while being exposed to blockers CNQX, AP5, picrotoxin and TTX. Subsequently, the membrane potential was stepped by depolarizing voltages, ranging from -60 mV to +40 mV in 10 mV increments, for a duration of 500 ms (Brereton et al., 2013). This protocol was employed to activate K+ currents (Figure 7A). First, we examined SK currents. The mean current density was determined at the end of the voltage pulses. In vehicle-treated, OVX females the application of the SK channel blocker apamin (100 nM) (Spergel, 2007) led to a significant reduction in whole-cell currents in the +10 to +40 mV range (Figures 7A, B). The mean outward current density at +40 mV in the control group was 125.5 ± 7.2 pA/pF (n = 5) with the apamin-sensitive component contributing 60.2 ± 3.3 pA/pF (n = 5) (Figures 7A, C). In contrast, in the E2-treated females, the overall mean outward current density at +40 mV was 195.4 ± 14.6 pA/pF (n = 6), which was significantly greater than the vehicle control group (Figures 7D, E). However, there was no significant difference in the apamin-sensitive component between the vehicle-treated and E2-treated females, 60.2 ± 3.3 pA/pF (n = 5) versus 75.9 ± 12.3 pA/pF (n = 6), respectively (Figure 7F). Furthermore, to investigate the expression of the mRNAs encoding SK channel subunits in Kiss1ARH neurons from vehicle-treated and E2-treated OVX females, qPCR experiments were performed on 10-cell Kiss1ARH neuronal pools (Figure 7G). We focused on SK3 channels because these channels exhibit the highest expression in the hypothalamus and E2 regulates their expression (Bosch et al., 2002). E2 treatment had no effect on the mRNA expression of the SK3 subunit. These molecular findings support our electrophysiology results. Our computational model was calibrated so that SK channels contributed ∼50 pA/pF to the whole-cell outward K+ current in E2-treated females (Figure 7H).

Contribution of Small conductance, calcium-activated K+ (SK) channel to Kiss1ARH neuronal excitability.

A. Representative traces of the inhibition of outward currents before (left, control) and after the specific SK blocker Apamin (500 nM, middle). Apamin sensitive currents were calculated from the subtraction of control and apamin at depolarized potentials (right). Cells were clamped at -70 mV and given 500 ms voltage pulses from -60 mV to +40 mV in 10 mV steps at 0.2 Hz, as shown in A at the bottom. B. Mean current density-voltage relationships measured at the end of the 500 ms voltage step ranging from -60 mV to +40 mV were obtained in the absence and presence of apamin (two-way ANOVA: main effect of treatment (F(1, 8) = 22.69, p = 0.0014), main effect of time (F(10, 80) = 306.0, p < 0.0001) and interaction (F(10, 80) = 24.76, p < 0.0001); mean ± SEM; n = 5; post hoc Bonferroni test, ***p < 0.005, ****p < 0.001). C. Apamin sensitive current densities were obtained from B (mean ± SEM, n = 5). D. Representative traces of the inhibition of outward currents before (left, control) and after the specific SK blocker Apamin (500 nM, middle). Apamin sensitive currents were resulted from the subtraction of control and apamin at depolarized potentials (right). E. Mean current density-voltage relationships measured at the end of the 500 ms voltage step ranging from -60 mV to +40 mV were obtained in the absence and presence of apamin (two-way ANOVA: main effect of treatment (F(1, 10) = 12.85, p = 0.0050), main effect of time (F(10,100) = 264.1, p < 0.0001) and interaction (F(10, 100) = 11.93, p < 0.0001); mean ± SEM, n = 6; post hoc Bonferroni test, **p < 0.01, ****p < 0.001). F. Apamin sensitive current densities were obtained from C and E (ns; Two-way ANOVA followed by Bonferroni post hoc test; mean ± SEM; OVX, n = 5; OVX+E2, n = 6). G. Kiss1ARH neurons (three 10-cell pools) were harvested from each of 5 vehicle- and 5 E2-treated, OVX females to quantify the mRNA expression of SK3 ion channel. E2 did not increase the mRNA expression of small conductance calcium-activated K+ (SK3) channels in Kiss1ARH. Bar graphs represent the mean ± SEM, with data points representing the number of animals, oil versus E2, Unpaired t test, t(8)= 0.551, p = 0.5967. The expression values were calculated via the ΔΔCT method, normalized to GAPDH and relative to the oil control values. H. The mathematical model was calibrated on the electrophysiology data from Kiss1ARH neurons before and after treatment with the specific SK blocker apamin, left panel versus middle panel respectively. For the calibration it was assumed that the applied concentration of apamin (500nM) completely blocked the SK current. The modeled apamin-sensitive current with gSK=28.1 nS (right panel) matches the electrophysiological data from OVX animals.

Additionally, following the same protocol and in the presence of the same cocktail of blockers (CNQX, AP5, picrotoxin and TTX), we investigated the contribution of BK channels to Kiss1ARH neuronal excitability. In the OVX females, the application of the BK channel blocker iberiotoxin (ibTx, 200 nM) (Niday and Bean, 2021) resulted in only a slight attenuation of the outward current (n = 5) (Figures 8A, B). The ibTx-sensitive current density measured at +40 mV was 32.8 ± 12.0 pA/pF (Figures 8A, C). However, in the E2-treated females, the application of ibTx significantly attenuated the whole-cell K+ current from +20 to +40 mV, (Figures 8D, E). Additionally, the ibTx-sensitive current was significantly larger in the +0 to +40 mV range in the E2-treated females compared to the OVX females (Figure 8F). These findings indicate that E2 treatment modulates the activity of ibTx-sensitive BK current in Kiss1ARH neurons, resulting in increased current density (90.0 ± 10.1 pA/pF (n = 6) versus 31.1 ± 8.4 pA/pF (n = 5) at +40 mV). To investigate the expression of mRNA encoding BK channel subunits in Kiss1ARH neurons from vehicle-treated and E2-treated OVX females, qPCR experiments were performed on 10-cell Kiss1ARH neuronal pools (Figure 8G). E2 treatment significantly increased the mRNA expression of the BKα1 (Kcnma1) subunit. These findings support our electrophysiological findings that there is a significant increase in BK channel activity in Kiss1ARH neurons with E2 treatment (Figure 8F). In addition, E2 increased the mRNA expression of Kcnb1 encoding delayed rectifier Kv2.1 (E2-treated relative mRNA expression: 1.672, n = 5 versus oil-treated mRNA expression: 1.086, n = 5; t-test p-value = 0.0024). The up-regulation of the two of these K+ channels (BKα1 and Kv2.1) would facilitate rapid repolarization of Kiss1ARH following an action potential. Our computational model was calibrated so that BK channels contributed ∼100 pA/pF to the whole cell outward K+ current in the E2-treated females (Figures 8H).

E2 up-regulates Kcnma1 mRNA and BK current in Kiss1ARH neurons.

Earlier version: A. Representative traces of the inhibition of outward currents before (left, control) and after the specific BK blocker iberiotoxin (IbTx; 200 nM, middle). IbTx sensitive currents were calculated from the subtraction of control and IbTx at depolarized potentials (right). Cells were clamped at -70 mV and given 500 ms voltage pulses from -60 mV to +40 mV in 10 mV steps at 0.2 Hz, as shown in A at the bottom. B. Mean current density-voltage relationships measured at the end of the 500 ms voltage step ranging from -60 mV to +40 mV were obtained in the absence and presence of IbTx (two-way ANOVA: main effect of treatment (F(1, 8) = 0.8841, p = 0.3746), main effect of time (F(10, 80) = 71.56), p < 0.0001) and interaction (F(10, 80) = 1.127, p = 0.3528); mean ± SEM, n = 5; post hoc Bonferroni test, p > 0.05). C. IbTX sensitive current densities were obtained from B (mean ± SEM, n = 5). D. Representative traces of the inhibition of outward currents before (left, control) and after the specific BK blocker iberiotoxin (IbTx; 200 nM, middle). IbTx sensitive currents were resulted from the subtraction of control and IbTx at depolarized potentials (right). E. Mean current density-voltage relationships measured at the end of the 500 ms voltage step ranging from -60 mV to +40 mV were obtained in the absence and presence of IbTX (two-way ANOVA: main effect of treatment (F(1, 10) = 4.660, p = 0.0562), main effect of time (F(10, 100) = 63.98, p < 0.0001) and interaction (F(10, 100) = 4.907, p < 0.0001); mean ± SEM, n = 6; post hoc Bonferroni test, *p < 0.05, **p < 0.01, **** < 0.001). F. IbTx sensitive current densities were obtained from C and E (two-way ANOVA: main effect of treatment (F(1, 9) = 22.04, p = 0.0011), main effect of time (F(10, 90) = 78.26, p < 0.0001) and interaction (F(10, 90) = 17.84, p <0.0001); mean ± SEM, OVX, n = 5; OVX+E2, n = 6; Bonferroni post hoc test, *p < 0.05, ***p < 0.005, ****p < 0.001). G. Kiss1ARH neurons (three 10-cell pools) were harvested from each of 5 vehicle- and 5 E2-treated, OVX females to quantify the mRNA expression of BKα channel. E2-treatment increased the mRNA expression of BKα. The expression values were calculated via the ΔΔCT method, normalized to GAPDH and relative to the oil control values. Bar graphs represent the mean ± SEM, with data points representing the number of animals (unpaired two-tailed t-test for BK, t(8) = 3.479, **p = 0.0083). H. The mathematical model was calibrated to reproduce the current voltage relationship observed in Kiss1ARH neurons before and after treatment with IbTx. For the calibration it was assumed that the applied concentration of IbTx (200 nM) completely blocked the BK current. The modeled IbTx -sensitive current with gBK=20.0 nS (right panel) matches the electrophysiological data from OVX+E2 animals.

Revised version: A. Representative traces of the inhibition of outward currents before (left, control) and after the specific BK blocker iberiotoxin (IbTx; 200 nM, middle). IbTx sensitive currents were calculated from the subtraction of control and IbTx at depolarized potentials (right). Cells were clamped at -70 mV and given 500 ms voltage pulses from -60 mV to +40 mV in 10 mV steps at 0.2 Hz, as shown in A at the bottom. B. Mean current density-voltage relationships measured at the end of the 500 ms voltage step ranging from -60 mV to +40 mV were obtained in the absence and presence of IbTx (two-way ANOVA: main effect of treatment (F(1, 8) = 0.8841, p = 0.3746), main effect of time (F(10, 80) = 71.56), p < 0.0001) and interaction (F(10, 80) = 1.127, p = 0.3528); mean ± SEM, n = 5; post hoc Bonferroni test, p > 0.05). C. IbTX sensitive current densities were obtained from B (mean ± SEM, n = 5). D. Representative traces of the inhibition of outward currents before (left, control) and after the specific BK blocker iberiotoxin (IbTx; 200 nM, middle). IbTx sensitive currents were resulted from the subtraction of control and IbTx at depolarized potentials (right). E. Mean current density-voltage relationships measured at the end of the 500 ms voltage step ranging from -60 mV to +40 mV were obtained in the absence and presence of IbTX (two-way ANOVA: main effect of treatment (F(1, 10) = 4.660, p = 0.0562), main effect of time (F(10, 100) = 63.98, p < 0.0001) and interaction (F(10, 100) = 4.907, p < 0.0001); mean ± SEM, n = 6; post hoc Bonferroni test, *p < 0.05, **p < 0.01, **** < 0.001). F. IbTx sensitive current densities were obtained from C and E (two-way ANOVA: main effect of treatment (F(1, 9) = 22.04, p = 0.0011), main effect of time (F(10, 90) = 78.26, p < 0.0001) and interaction (F(10, 90) = 17.84, p <0.0001); mean ± SEM, OVX, n = 5; OVX+E2, n = 6; Bonferroni post hoc test, *p < 0.05, ***p < 0.005, ****p < 0.001). G. Kiss1ARH neurons (three 10-cell pools) were harvested from each of 5 vehicle- and 5 E2-treated, OVX females to quantify the mRNA expression of BKα channel. E2-treatment increased the mRNA expression of BKα. The expression values were calculated via the ΔΔCT method, normalized to GAPDH and relative to the oil control values. Bar graphs represent the mean ± SEM, with data points representing the number of animals (unpaired two-tailed t-test for BK, t(8) = 3.479, **p = 0.0083). H. The mathematical model was calibrated to reproduce the current voltage relationship observed in Kiss1ARH neurons before and after treatment with IbTx. For the calibration it was assumed that the applied concentration of IbTx (200nM) completely blocked the BK current. The modeled IbTx -sensitive current with gBK=20.0 nS (right panel) matches the electrophysiological data from OVX+E2 animals.

Traditionally, the after hyperpolarization is divided into three distinct phases: fast (fAHP), medium (mAHP), and slow after hyperpolarization (sAHP) (Storm, 1990; Vogalis et al., 2003). The fast after hyperpolarization (fAHP) is primarily mediated by the BK family of potassium channels (Storm, 1987). The medium after hyperpolarization (mAHP) is predominantly mediated by apamin-sensitive SK2 channels (Bond et al., 2004; Peters et al., 2005). However, KCNQ family members contribute to both the mAHP and sAHP (Tzingounis A. V. et al., 2010; Tzingounis A. V. and Nicoll, 2008). Therefore, to investigate the contribution of KCNQ channels to Kiss1ARH neuronal excitability, voltage clamp experiments were conducted in the presence of TTX, CNQX, AP5, and picrotoxin, and a standard M-current protocol was run using the M-channel blocker XE-991 to isolate the M-current (Greene et al., 2017; Roepke et al., 2011) (Figure 9 A, B). The application of XE-991 resulted in the inhibition of M-current within a physiologically relevant voltage range of -60 to -30 mV in E2-treated OVX females but exhibited minimal impact in OVX females (Figure 9 C, D). Although the XE991-sensitive current was relatively small compared to other voltage-activated K+ conductances, it demonstrated a significant increase in E2-treated, OVX females (Figures 9 E). The maximum peak current density sensitive to XE-991 at -30 mV was found to be four times higher in E2-treated OVX females when compared to OVX females. This would contribute to the repolarization following burst firing. Furthermore, E2 increased the mRNA expression of Kcnq2, (Figure 9F), which suggests that KCNQ channels play a role in repolarizing Kiss1ARH neurons following burst firing. Indeed, our modeling predicted that M-current contributed to the repolarization following burst firing (Figure 9G).

E2 up-regulates KCNQ channels and M-current in Kiss1ARH neurons.

A-B. Representative current traces of the M-current inhibition caused by 40 µM XE-991 perfused for 10 min in (A) OVX-oil and (B) OVX+E2-treated female mice. Inset: M-current deactivation protocol. C-D. Current density-voltage plots from – 75 to –30 mV of vehicle and XE-991 perfusion in (C) OVX-oil and (D) OVX+E2-treated mice. Two-way ANOVA for C: main effect of treatment (F(1, 17) = 1.908, p = 0.1851), main effect of time (F(9, 153) = 187.1, p < 0.0001), and interaction (F(9, 153) = 3.901, p = 0.0002); Veh, n = 11; XE-991, n = 8; Bonferroni post hoc test, p > 0.05. For D: main effect of Veh and XE-991 (F(1, 24) = 24.92, p < 0.0001), main effect of time (F(9, 216) = 174.5, p < 0.0001), and interaction (F(9, 216) = 52.75, p < 0.0001); Veh, n = 13; XE-991, n = 13; Bonferroni post hoc test, a = p < 0.05, b = p < 0.001. E. Treatment with E2 elevated, while XE-991 diminished the maximum peak current density elicited by a -30 mV step in OVX- and OVX+E2-treated mice. Two-way ANOVA: main effect of Veh and XE-991 (F(1, 41) = 47.59, p < 0.0001), main effect of OVX and OVX+E2 (F(1, 41) = 15.76, p = 0.0003), and interaction (F(1, 41) = 18.2, p = 0.0001; Veh, n = 11; XE-991, n = 8; Bonferroni post hoc test, Veh: OVX vs. OVX+E2, a = p < 0.001. XE-991: OVX vs. OVX+E2, p > 0.05. Data are expressed as mean ± SEM, with data points representing the number of cells. F. Kiss1ARH neurons (three10-cell pools) were harvested from each of 5 vehicle- and 5 E2-treated, OVX females to quantify the mRNA expression of Kcnq2. E2 treatment increased the mRNA expression of Kcnq2. Unpaired t-test, t(8) = 4.850, **p = 0.0013. Data are expressed as mean ± SEM, with data points representing the number of animals. G. Percent contribution of the different K+ currents to the repolarization current during burst-type firing activity in the OVX+E2 state. At each time point, the length of each color bar denotes the percent contribution of the corresponding current to the total outward current.

E2 increases Vglut2 but down regulates Tac2, Trpc5 and Girk2 mRNA expression in Kiss1ARH neurons

Based on our electrophysiological results, Kiss1ARH neurons appear to transition from peptidergic to glutamatergic neurotransmission through E2-mediated changes in the expression of voltage-activated Ca2+ channels and K+ channels and their respective conductances. Therefore, we asked the question is there a difference in peptide and glutamate mRNA expression mediating this transition? Therefore, we ran a comparison between Tac 2 (NKB) and Vglut2 (surrogate for glutamate) expression. The cycle threshold (CT) was compared between Tac2 and Vglut2 as well as Kiss1, TRPC5 and GIRK2 in Kiss neuronal cell pools from OVX oil-treated and OVX E2-treated animals (Figures 10A, B). It is worth noting that lower number of cycles illustrate a higher quantity of mRNA expression because the fluorescence is detected earlier, and one cycle difference represents a doubling in expression. As expected, the reference gene Gapdh did not change with E2-treatment. However, quantitative PCR results revealed that E2 treatment of OVX females significantly reduced Tac2 expression (Figure 10C), whereas Vglut2 mRNA was significantly increased in Kiss1ARH neurons (Figure 10D). Moreover, both Trpc5 and Girk2 expression were significantly reduced in E2-treated, OVX females (Figures 10E, F).

Estradiol decreases Tac 2, Trpc5 and Kcnj6 but increases Vglut 2 mRNA expression in Kiss1ARH neurons.

Earlier version: A. qPCR amplification curves illustrating the cycle threshold (CT) for Tac2, Gapdh, Kiss1, Trpc5, Slc17a6 (Vglut2) and Kcnj6 (GIRK2) in Kiss1ARH ten cell neuronal pools (three to six pools from each animal) in OVX Oil-treated, and B, in OVX E2-treated females. C. Quantitative real-time PCR analysis of Tac2 mRNA, D, Vglut2, E, Trpc5, F, Kcnj6. Comparisons were made between Oil-treated and E2-treated, OVX females using the comparative 2-ΔΔCT method. Bar graphs represent the mean ± SEM, with data points representing the number of animals. (Unpaired t test for Tac2, t(6) = 10.670, ***p < 0.0001; for Vglut2, t(7) = 9.678, ***p < 0.0001; for Trpc5, t(6) = 5.774, **p = 0.0012; for Kcnj6, t(8) = 3.457, **p = 0.0086).

Revised version: A qPCR amplification curves illustrating the cycle threshold (CT) for Tac2, Gapdh, Kiss1, Trpc5, Slc17a6 (Vglut2) and Kcnj6 (GIRK2) in Kiss1ARH ten cell neuronal pools (three to six pools from each animal) in OVX Oil-treated, and B, in OVX E2-treated females. C. Quantitative real-time PCR analysis of Tac2 mRNA, D, Vglut2, E, Trpc5, F, Kcnj6. Comparisons were made between Oil-treated and E2-treated, OVX females using the comparative 2-ΔΔCT method. Bar graphs represent the mean ± SEM, with data points representing the number of animals. (Unpaired t test for Tac2, t(6) = 10.670, ***p < 0.0001; for Vglut2, t(7) = 9.678, ***p < 0.0001; for Trpc5, t(6) = 5.774, **p = 0.0012; Unpaired t test for Kcnj6, t(8) = 3.457, **p = 0.0086).

CRISPR mutagenesis of Trpc5 attenuates slow EPSP and reduces excitability of Kiss1ARH neurons

Earlier version: Our previous (Qiu Jian et al., 2021) and current findings suggested that TRPC5 channels play a dominant role in regulating cell excitability in ovariectomized females. Therefore, as proof of principle, we utilized a CRISPR approach to mutate TRPC5 channels in Kiss1ARH neurons similar to our previous studies (Hunker et al., 2020; Stincic et al., 2021). Hunker et al. developed a single viral vector for conditional expression of the smaller Staphylococcus aureus (SaCas9) and sgRNA that yields high-efficiency mutagenesis in specific cell types (Hunker et al., 2020). To selectively mutate Trpc5 in Kiss1ARH neurons, we generated two guide RNA’s, one targeting exon 2, which is conserved across all splice variants, and the other targeting exon 7, the pore forming domain (Figures 11A, B). A cohort of Kiss1ARH mice were given bilateral stereotaxic injections into the ARH of the two AAV1-FLEX-SaCas9-sgTrpc5’s or a control virus containing with AAV1-FLEX-SaCas9-U6-sgRosa26. An additional Cre-dependent virus of the same serotype (AAV1) that drove expression of a fluorophore (YFP or mCherry) was co-administered in order to visualize injection quality and facilitate harvesting of cells (Figure 11C). After three weeks, mice underwent ovariectomy since OVX mice express the maximum slow EPSP amplitude (Qiu J. et al., 2016). Brain slices were prepared, and cells harvested as previously described (Qiu J. et al., 2018) and analyzed with qPCR. We found that the Trpc5 mutagenesis group displayed a reduction in relative expression of Trpc5 in Kiss1ARH neurons compared to the control group (Figure 11D). Hence, the qPCR data verified that in the sgTrpc5-targeted mice we can selectively reduce Trpc5 gene expression in targeted cells.

Revised version: Our previous (Qiu Jian et al., 2021) and current findings suggested that TRPC5 channels play a dominant role in regulating cell excitability in ovariectomized females. Therefore, as proof of principle, we utilized a CRISPR approach to mutate TRPC5 channels in Kiss1ARH neurons similar to our previous studies (Hunker et al., 2020; Stincic et al., 2021). Hunker et al. developed a single viral vector for conditional expression of the smaller Staphylococcus aureus (SaCas9) and sgRNA that yields high-efficiency mutagenesis in specific cell types (Hunker et al., 2020). To selectively mutate Trpc5 in Kiss1ARH neurons, we generated two guide RNA’s, one targeting exon 2, which is conserved across all splice variants, and the other targeting exon 7, the pore forming domain (Figures 11A, B). A cohort of Kiss1ARH mice were given bilateral stereotaxic injections into the ARH of the two AAV1-FLEX-SaCas9-sgTrpc5’s or a control virus containing with AAV1-FLEX-SaCas9-U6-sgRosa26.

CRISPR mutagenesis of Trpc5 channels in Kiss1ARH neurons.

A, structure of AAV1-FLEX-SaCas9-U6sgTrpc5-exon2. Exon 2 of Trpc5 is denoted with guide sequence highlighted in red, the PAM is underlined. B. Structure of AAV1-FLEX-SaCas9-U6sgTrpc5-exon7. Exon 7 of Trpc5 is denoted with guide sequence highlighted in red, the PAM is underlined. C1, image of coronal section through the ARH from Kiss1-Cre::Ai32 mouse with dual co-injections of AAV-DIO-mCherry and AAV1-FLEX-SaCas9-U6-sgTrpc5. Scale = 200 µm. C2, C3, higher power overlays of epifluorescence (EYFP & mCherry) images with recording pipette patched onto a Kiss1ARH-Cre:mCherry cell (C3). Scale = 40 µm. D, quantitative PCR measurements of Trpc5 transcripts from 10 cell neuronal pools (three pools from each animal) in double sgRNA mutagenesis of Trpc5 (second sgRNA against pore forming region) in Kiss1ARH neurons. Primers were targeted to 1st (left panel) or 2nd (right panel) guide, respectively (Unpaired t test for 1st, t(10) = 10.67, ****p < 0.0001; for 2nd, t(10) = 10.79, ****p < 0.0001). Data are expressed as mean ± SEM, with data points representing the number of animals.

Added in the revised version: An additional Cre-dependent virus of the same serotype (AAV1) that drove expression of a fluorophore (YFP or mCherry) was co-administered in order to visualize injection quality and facilitate harvesting of cells (Figure 11C). After three weeks, mice underwent ovariectomy since OVX mice express the maximum slow EPSP amplitude (Qiu J. et al., 2016). Brain slices were prepared, and cells harvested as previously described (Qiu J. et al., 2018) and analyzed with qPCR. We found that the Trpc5 mutagenesis group displayed a reduction in relative expression of Trpc5 in Kiss1ARH neurons compared to the control group (Figure 11D). Hence, the qPCR data verified that in the sgTrpc5-targeted mice we can selectively reduce Trpc5 gene expression in targeted cells.

As predicted, mutagenesis of Trpc5 channels in Kiss1ARH neurons significantly attenuated the slow EPSP (Figures 12A, B, C) such that the postsynaptic excitation was reduced to a “trickle” of action potential firing. What we would not have predicted is that the double sgRNA mutagenesis of Trpc5 channels in Kiss1ARH neurons significantly hyperpolarized the resting membrane potential by 7 mV (Figure 12D). Moreover, the rheobase (minimum current required to induce firing) significantly increased by ∼13 % in females bearing the sgTrpc5 double mutagenesis (Figure 12E). The firing frequency versus injected current (F-I) curve for sgTrpc5 double mutagenesis Kiss1ARH neurons was also significantly attenuated (Figure 12F). In agreement with these experimental findings, current ramp simulations of our mathematical model confirmed that a decrease of the TPRC5 conductance increased the rheobase and reduce the firing activity of Kiss1ARH neurons (Figures 12G, H). In our model, this contribution of TRPC5 channels to cell excitability is attributed to intracellular calcium levels which activate TRPC5 channels independently of NKB stimulation. Finally, we employed our mathematical model to further investigate the transition from synchronous firing driven by NKB release and TRPC5 channel activation to burst firing generated by E2-mediated upregulation of endogenous conductances. As we varied the conductance of TRCP5 channels and GIRK channels under saturating extracellular signals (e.g., NKB and Dyn), we find that balance between the two conductances controls the neuronal excitability in the OVX parameterized model, with synchronous firing eliminated when TRPC5 conductance is low (red triangle) relative to the GIRK conductance (Figure 13A). Furthermore, the burst firing of the OVX+E2 parameterized model was supported by elevated h- and Ca 2+-currents (Figure 13B) as well as by the high conductance of Ca2+ channels relative to the conductance of TRPC5 channels (Figure 13C). Importantly, the simulated burst firing was analogous to the phasic burst firing recorded ex vivo in Kiss1ARH neurons (Figure 1A).

Double CRISPR mutagenesis of Trpc5 attenuates slow EPSP, increases rheobase and shifts the F-I curve.

A, high-frequency photo-stimulation (20 Hz) generated slow EPSP in Kiss1ARH neuron from ovariectomized, control mouse. Red trace is slow EPSP after low-pass filtering. B, slow EPSP in Kiss1ARH neuron from OVX, double sgTrpc5 -targeted (Trpc5 mut) mouse. C, summary of the effects of Trpc5 mutagenesis on slow EPSP amplitude in female mice. Unpaired t-test, t(27) = 5.916, ****p < 0.0001. Data are expressed as mean ± SEM, with data points representing the number of cells. D, double sgRNA mutagenesis of Trpc5 channels in Kiss1ARH neurons significantly increased the RMP (control: -64.5 ± 1.4 mV versus Trpc5 mut, 2: -71.1 ± 1.2 mV, Unpaired t-test, t(73) = 3.524, ***p = 0.0007). E, current ramp showing the increased rheobase in Kiss1ARH neuron from Trpc5 mut mice (control: 31.1 ± 1.2 pA, n = 31, versus Trpc5 mut, 35.3 ± 1.0 pA, n = 33, Unpaired t-test, t(62) = 2.777, **p = 0.0073). F, firing frequency vs. current (F-I) curves for control versus Trpc5 mut (two-way ANOVA: main effect of treatment, F(1,52) = 13.03, p = 0.0007, main effect of time, F(8,416) = 291.3, p < 0.0001, and interaction, F(8,416) = 6.254, p < 0.0001; Control, n = 26, Trpc5 mut, n = 29; post hoc Bonferroni test, *<0.05; **<0.01 and ***p < 0.005, respectively). G, model simulations of the effects of a current ramp (50 pA/sec) for OVX (left panel) and OVX female with reduced (muted) TRPC5 conductance (right panel), and H, the associated firing frequency vs. current curves. In the latter case the TRCP5 conductance was halved, which is a conservative estimation of the CRISPR state in which the Trpc5 is much more mutated in Kiss1ARH neurons (Figure 11).

Computational modeling of a Kiss1ARH neuron in the OVX and OVX + E2 state demonstrates its distinct dynamic responses.

A model of the Kiss1ARH neuron was developed and calibrated using molecular data and electrophysiological recordings of Kiss1ARH neurons from OVX and OVX+E2 mice. A. Simulations of the OVX-parameterized model demonstrating high frequency activity in response to saturating levels of NKB stimulation. The balance between GIRK and TRCP5 conductance controls the response of the neuron to NKB stimulation, with neuronal response eliminated when TRPC5 conductance is low (red triangle) relative to the GIRK conductance. B. The OVX+E2 parameterized models demonstrate sustained burst firing activity. The bursting activity that is supported by elevated h- and Ca 2+-currents (red square) as observed in OVX+E2 mice. C. In the OVX+E2 state, burst firing activity is also supported by high conductance of HVA Ca2+ channels relative to the conductance of TRPC5 channels. Representative points in the parameter space giving rise to burst firing activity are marked with red squares, whereas red triangles are used for points resulting in regular spiking. The black line separates these two regions of activity.

Discussion

E2 appears to play a critical role in transitioning the glutamatergic/peptidergic Kiss1ARH neurons from a high frequency firing mode during synchronization, which is dependent on NKB-driven activation of TRPC5 channels, to a short bursting mode that would facilitate glutamate release. E2 decreased the expression of the peptide neurotransmitters NKB (kisspeptin and dynorphin) and TRPC5 channels but increased the mRNA expression of Vglut2, voltage-activated calcium channels and Hcn channels that contribute to burst firing and glutamate release from the Kiss1ARH neurons. We determined that the increase in mRNA expression of the HVA calcium channels translated into a significant increase in whole-cell current with all of the calcium channels contributing proportionally. Most importantly the kinetics of activation and inactivation were unaltered with E2 treatment, which indicates that other post-translation modifications were not affecting channel activity. Surprisingly and somewhat counter intuitive, the BK α1 subunit was also upregulated, but based on our modeling the rapid repolarization of the Kiss1ARH neurons (i.e., the fast AHP) would facilitate higher frequency of action potential firing. Moreover, our modeling confirmed that TRPC5 channels, which generate the slow EPSP (a.k.a., plateau potential in other CNS neurons), are vital for initiating and sustaining synchronous firing of Kiss1ARH neurons, while subsequent activation of GIRK channels repolarizes Kiss1ARH neurons. E2 treatment of ovariectomized females decreased both Trpc5 and Girk2 channel mRNA expression, which in our model correlated with the reduction in sustained high frequency firing of Kiss1ARH neurons. Therefore, the synchronous high frequency firing of Kiss1ARH neurons in a low E2 milieu correlates with the pulsatile release of GnRH (LH from the pituitary gland), whereas the transition to burst firing in the presence of high circulating levels of E2 (e.g., proestrus) would facilitate the GnRH (LH) surge through its glutamatergic synaptic connection with Kiss1AVPV/PeN neurons (Lin et al., 2021; Qiu J. et al., 2016).

Core calcium conductances underlying synchronous and burst firing of Kiss1ARH neurons

TRPC5 channels are highly expressed in Kiss1ARH neurons (Figure 11), and TRPC5 channels are essentially ligand-activated calcium channels with a high permeability to calcium (PCa/PNa = 9:1) (Venkatachalam and Montell, 2007). In general, mammalian TRPC channels are activated by both G protein-coupled receptors and receptor tyrosine kinases (Ambudkar and Ong, 2007; Clapham, 2003), and are one of the major downstream effectors activated by glutamate binding to group I metabotropic glutamate receptors (mGluR1 and mGluR5) in CNS neurons (Bengtson et al., 2004; Berg et al., 2007; Faber et al., 2006; Tozzi et al., 2003). In substantia nigra dopamine neurons mGluR1 agonists induce a current that exhibits the tell-tale double-rectifying current-voltage plot of TRPC channel activation (Tozzi et al., 2003), similar to what we see with the effects of the NKB agonist senktide in Kiss1ARH neurons (Qiu Jian et al., 2021). Both mGluR1 and TacR3 are Gq-coupled to phospholipase C (PLC) activation which leads to hydrolysis of phosphatidylinositol 4,5-bisphosphate (PIP2) to diacylglycerol (DAG) and inositol 1,4,5 triphosphate (IP3), which is involved in channel activation (Birnbaumer, 2009). TacR3 (NKB) signaling has additional consequences since many K+ (e.g., GIRK, KCNQ) channels are dependent on PIP2 for channel opening, and depletion of PIP2 by PLC leads to channel closure (Brown and Passmore, 2009) (Zhang C. et al., 2013) (Whorton and MacKinnon, 2011) (Zheng et al., 2022). Therefore, depletion of PIP2 by NKB signaling would further facilitate the sustained firing of Kiss1ARH neurons during synchronization.

A plateau potential has been characterized in hippocampal and cortical neurons (Arboit et al., 2020; Zhang Z. et al., 2011) as such these neurons express biophysical properties that allow them to continue to persistently fire even after a triggering synaptic event has subsided (Zylberberg and Strowbridge, 2017). The persistent firing activity of these neurons is linked to ICAN (Zylberberg and Strowbridge, 2017), and TRPC5 channels appear to be responsible for the ICAN (Zhang Z. et al., 2011). With TacR3 activation in Kiss1ARH neurons there is an influx of Ca+2 through TRPC5 channels leading to greater build-up of [Ca2+]i that facilitates the opening of more TRPC5 channels in a self-sustaining (autocatalytic) manner (Qiu J. et al., 2016). Using the fast intracellular calcium chelator BAPTA, which has been shown to robustly inhibit TRPC5 channel activation in heterologous cells (Blair et al., 2009), we abolished the slow EPSP and persistent firing of Kiss1ARH neurons following optogenetic stimulation (Qiu Jian et al., 2021). Moreover, HVA calcium channel blockers attenuated the generation of the slow EPSP (Figure 3) so it appears that they also contribute to the ICAN since calcium influx via both LVA and HVA calcium channels can also facilitate TRPC5 channel opening in Kiss1ARH neurons. In the ovariectomized female, treatment with E2 upregulated Cav1.2, Cav2.1, Cav2.2, and Cav2.3 mRNA by 1.5-to 2-fold and Cav3.1 mRNA expression by ∼3-fold. Hence, E2 significantly increased whole-cell calcium currents Kiss1ARH neurons, which greatly enhanced the excitability and contributed to the burst firing of Kiss1ARH neurons (Figure 1). However, the amplitude of the slow EPSP with E2 treatment is only ∼25% of the amplitude in the ovariectomized state (Qiu J. et al., 2016). Therefore, there appears to be a physiologic transition of Kiss1ARH neurons from the slow EPSP firing mode in the OVX state to the burst firing mode in the presence of E2, which has important physiologic ramifications as discussed below.

TRPC5 and GIRK channels are vital for synchronization of Kiss1ARH neurons

Recently, Tian et al. demonstrated that TRPC4, a close homolog of TRPC5 sharing ∼64 percent homology, is a “coincidence detector” of neurotransmission by both Gq/11 and Gi/o-coupled receptors in lateral septal (LS) neurons (Tian et al., 2022). In whole-cell recordings of LS neurons, TRPC 4 channels mediate a strong depolarizing plateau potential that in contrast to TRPC5 channel activation in Kiss1ARH neurons, abrogates action potential firing as a result of a depolarization block. In many instances the plateau potential in LS neurons is followed by an AHP, which is dependent on the activation of Gi/o-coupled receptors. In contrast, we have not observed an AHP in Kiss1ARH neurons following the slow EPSP (Qiu J. et al., 2016). Tian and colleagues showed that the depolarizing plateau in LS neurons is codependent on activation of both Gq/11-coupled mGluR1 glutamate receptors and Gi/o-coupled γ-aminobutyric acid type B receptors, the latter activating GIRK channels. Moreover, the firing patterns in LS neurons encodes information about the relative strengths of these contrasting inputs (i.e., Gq/11 versus Gi/o) such that only mGluR1 produces weak depolarization accompanied by increased firing of LS neurons, whereas pure GABAB receptor activation hyperpolarizes the cells and abrogates firing activity. Coincident input of both mGluR1 and GABAB receptors results in a brief burst of action potentials followed by a pause in firing, and both the pause duration and firing recovery patterns reflect the relative strengths of Gq/11 versus Gi/o inputs. Importantly, Tian and colleagues computationally simulated these various scenarios with computational modeling, and similar to our modeling, utilized only TRPC4 and GIRK channels. A notable difference between the Kiss1ARH neurons and the LS neurons is that dynorphin binds to kappa-opioid receptors to open GIRK channels in the nerve terminal of Kiss1ARH neurons (presynaptic effect) (Qiu J. et al., 2016), whereas GABA binds to some GABAB receptors to hyperpolarize LS neurons (Tian et al., 2022). Therefore, although the high frequency firing activity of both LS and Kiss1ARH neurons can be modeled around TRPC and GIRK channels, the generated firing patterns are dramatically different based on the timing (co-incident activation of TRPC4 and GIRK channels in LS) and localization of the kappa opioid-coupled GIRK channels in the axon terminal of Kiss1ARH neurons. Moreover, E2-treated, ovariectomized females show a significant down-regulation of both Trpc5 and Girk2 mRNA expression in Kiss1ARH neurons (Figure 10), which is important for the physiological transitioning as described below.

Interestingly, the hypothalamic A12 (ARH) dopamine neurons show a rhythmic “oscillatory” firing behavior that transitions to a tonic firing mode with synaptic input from Thyrotropin-releasing hormone (TRH) neurons (Lyons D.J. et al., 2010) or feedback by circulating prolactin, released by pituitary lactotrophs (Lyons D. J. et al., 2012). A more recent paper has revealed that TRPC5 channels mediate the plateau potential and tonic firing in A12 dopamine neurons in response to prolactin (Blum et al., 2019). Similar to our findings with CRISPR deletion of Trpc5 in Kiss1ARH neurons (Figures 11 and 12), conditional knockout of Trpc5 in dopamine neurons abrogated the prolactin-induced plateau potential and tonic firing. Although Blum and colleagues did not model the oscillatory firing or tonic firing of the A12 dopamine neurons, their findings are consistent with our results showing that the activation of TRPC5 channels underlies the slow EPSP (plateau potential) and sustained firing.

Contribution of endogenous K+ channels to synchronized and burst firing

Beyond the ligand-gated (e.g., baclofen) GIRK channels, there are endogenous K+ channels that help sculpt the firing activity of kisspeptin neurons. We focused on the calcium-activated K+ channel family: the large-conductance, calcium-activated potassium (BK, also called BKCa, KCa1.1, MaxiK, Slo), small conductance, calcium-activated K+ (SK1, SK2, SK3) (Bond et al., 2005), and the K+ channels underlying the M-current (KCNQ, Kv7.1-7.5) (Brown and Passmore, 2009), which mediate the fast afterhyperpolarization (AHP), the intermediate AHP/slow AHP, respectively (Andrade et al., 2012) BK channels are gated by both voltage and cytoplasmic calcium and sculpt action potential firing in CNS neurons (Blatz and Magleby, 1987; Marty, 1989; Sah P., 1996; Storm, 1990). Indeed, BK channels have been shown to mediate rapid spike repolarization—i.e., the fast AHP in hippocampal CA1 pyramidal neurons (Lancaster and Nicoll, 1987; Storm, 1987). Blockade of BK channels in CA1 neurons attenuates the initial discharge frequency in response to current injection, which is attributable to suppression of the BK channel-dependent rapid spike repolarization (Lancaster and Nicoll, 1987; Storm, 1987). Blockade of BK channels is thought to increase inactivation of the spike-generating transient Na+ current and activate more of the slower K+ currents, thereby enhancing refractoriness and reducing excitability during the immediate aftermath of the first action potential (Shao et al., 1999). Thus, BK channels facilitate high-frequency burst firing of CA1 neurons. Furthermore, extracellular field recordings confirmed that BK channels contribute to high-frequency burst firing in response to excitatory synaptic input to distal dendrites in CA1 neurons (Gu et al., 2007). Therefore, BK channels appear to play an important role for early high-frequency, rapidly adapting firing in hippocampal CA1 pyramidal neurons, thus promoting the type of bursting that is characteristic of these cells in vivo during behavior. Based on our in vitro electrophysiological recordings and computational modeling we see a similar physiological phenomenon in Kiss1ARH neurons (Figure 8). Not only does E2 increase the mRNA expression of Kcnma1 (Figure 8G), but also the maximum BK (IbTx-sensitive) current by 3-fold, which would facilitate rapid repolarization during burst firing and promote glutamate release similar to hippocampal CA1 neurons.

In contrast to BK channel expression, E2 did not affect the mRNA expression of SK3 channel mRNA. SK channels underlie the apamin-sensitive component of the medium duration AHP and are responsible for repolarization following a burst of action potentials (Andrade et al., 2012; Bond et al., 2005). The activation of SK channels is voltage-independent, but SK channels have a higher affinity for Ca2+ than BK channels (Andrade et al., 2012). SK channels are tightly coupled (within 100 nm) to L-type Ca2+ channels, and BK channels (within 30 nm) are tightly coupled to N-type Ca2+ channels in hippocampal CA1 pyramidal neurons (Marrion and Tavalin, 1998). The determination of the proximity of SK and BK channels to HVA calcium channels in kisspeptin neurons will require cell-attached patch recordings. However, in Kiss1ARH neurons the SK channels may come into play during a short burst of action potentials but would become overwhelmed with the higher frequency synchronized, sustained firing as a result of NKB stimulation (Qiu J. et al., 2016). As discussed above what limits the synchronized firing of Kiss1ARH neurons is the activation of GIRK channels.

Finally, the calcium-activated slow AHP probably plays a critical role in the repolarization of Kiss1ARH after burst firing. The molecular identification of the channels mediating the slow AHP has long been an area of intense investigation (Andrade et al., 2012; Vogalis et al., 2003). A critical feature of the slow AHP is that it activates very slowly (hundreds of milliseconds) long after the rise in cytoplasmic Ca2+ (Sah P. and Clements, 1999) so an intermediate Ca2+ signaling molecule has long been thought to be involved. Indeed, Tzingounis and colleagues (Tzingounis A.V. et al., 2007) have provided compelling evidence that the diffusible calcium sensor hippocalcin is the critical intermediate molecule involved in Ca2+ sensing. The slow AHP is abrogated in hippocalcin KO mice, and transfection of hippocalcin into cultured hippocampal neurons generates a pronounced slow AHP in response to a depolarizing stimulus (Tzingounis A.V. et al., 2007). Importantly, the slow AHP is activated by Ca2+ with an EC50 ≈ 300 nM, which is well within the operational range of hippocalcin but well below that of calmodulin (Andrade et al., 2012). Finally, two seminal papers from Tzingounis and colleagues demonstrate that KCNQ 2, 3 channels are responsible for the slow AHP in hippocampal dentate neurons (Tzingounis A. V. and Nicoll, 2008), and KCNQ 5 channels are responsible for the slow AHP in CA3 neurons (Tzingounis A. V. et al., 2010). Moreover, the KCNQ channel blocker XE991 attenuates the slow AHP in CA3 neurons (Tzingounis A. V. and Nicoll, 2008). Based on these seminal findings we investigated the role of the KCNQ channels, which “classically” underlie the M-current in Kiss1ARH. The M-current was first identified in Kiss1ARH neurons by Conde and Roepke (Conde and Roepke, 2019), and presently we found that Kcnq2 mRNA is expressed in Kiss1ARH neurons and up-regulated by E2, which translated into a greater M-current in Kiss1ARH neurons (Figure 9). Incorporating the M-current into our computational model indeed supports our hypothesis that M-current, along with ISK and IBK, contributes to membrane repolarization after burst firing (Figure 9G). Importantly the slow AHP, as opposed to the fast AHP (BK) and medium AHP (SK) is highly regulated by multiple neurotransmitters (Andrade et al., 2012), which sets the stage for further modulation of the slow EPSP in Kiss1ARH neurons.

Importance of E2-driven physiological transitioning

Since the expression of the peptide neurotransmitters in Kiss1ARH neurons are down-regulated by E2, the Kiss1ARH neurons are believed to be under “inhibitory” control by E2 and are important for “negative-feedback” regulation of GnRH release (Lehman et al., 2013; Navarro V. M. et al., 2009; Smith et al., 2005) (Rance Naomi E. and Young, 1991) (Rance N.E., 2009). However, our past (Gottsch et al., 2011) and current findings document that these Kiss1ARH neurons express HVA and LVA calcium and HCN (pacemaker) channels and are excited by co-released glutamate from neighboring Kiss1ARH neurons, which indicates that these neurons have pacemaker electrophysiological properties similar to other CNS neurons (Bal and McCormick, 1993; Lüthi and McCormick, 1998). Additionally, in contrast to the neuropeptides, E2 increases Slc17a6 (Vglut2) mRNA expression, in addition to mRNA for the voltage-activated calcium and HCN channels, and increases glutamate release onto Kiss1AVPV/PeN neurons (Qiu J. et al., 2018). Interestingly, Slc17a6 mRNA expression in Kiss1ARH neurons and the probability of glutamate release are decreased along with the neuropeptides in intact versus castrated males (Nestor et al., 2016), which indicates a profound sex difference in the glutamate signaling by Kiss1ARH neurons (Nestor et al., 2016; Qiu J. et al., 2018). Obviously, in the male there is no preovulatory LH surge so there is no need for excitatory glutamatergic input to the few Kiss1AVPV neurons in the male.

In females, conditional knockout of Slc17a6 in Kiss1 neurons abrogates glutamate release from Kiss1ARH neurons (Qiu J. et al., 2018). Kiss1AVPV/PeN neurons do not express Slc17a6 and do not release glutamate. Within the Kiss1ARH neurocircuitry the lack of glutamate transmission does not diminish the slow EPSP in ovariectomized females (Qiu J. et al., 2018). Indeed, a recent publication from the Herbison lab demonstrates that glutamate generates small “synchronizing” events that are dependent on the ionotropic receptors (Han et al., 2023), but the fast neurotransmitter is unable to support the sustained firing (i.e., slow EPSP) that is necessary for peptide release and synchronization of the KNDy network. Rather, we postulate that glutamate neurotransmission is more important for excitation of Kiss1AVPV/PeN neurons and facilitating the GnRH (LH) surge with high circulating levels of E2 when peptide neurotransmitters are at a nadir and glutamate levels are high in female Kiss1ARH neurons. Indeed, low frequency (5 Hz) optogenetic stimulation of Kiss1ARH neurons, which only releases glutamate in E2-treated, ovariectomized females (Qiu J. et al., 2016), generates a surge-like increase in LH release during periods of optical stimulation (Lin et al., 2021; Voliotis et al., 2021). In a subsequent study optical stimulation of Kiss1ARH neuron terminals in the AVPV at 20 Hz, a frequency commonly used for terminal stimulation in vivo, generated a similar surge of LH (Shen et al., 2022). Additionally, intra-AVPV infusion of glutamate antagonists, AP5+CNQX, completely blocked the LH surge induced by Kiss1ARH terminal photostimulation in the AVPV (Shen et al., 2022). Therefore, there appears to be a clear role for glutamatergic transmission from the Kiss1ARH to Kiss1AVPV/PeN neurons in amplifying the LH surge in the female mouse. Finally it is important to keep in mind that even in the presence of high physiological levels of E2, the mRNA expression of Tac2 is many-fold higher than Kiss1 (Figure 10), which is essential for NKB maintaining synchronous firing of Kiss1ARH neurons, albeit at a lower frequency, across all physiological states (Qiu J. et al., 2016). Indeed, there is a progressive change from a strictly pulsatile pattern of GnRH in the hypophyseal portal circulation to one containing both pulsatile and non-pulsatile components during the development of the GnRH surge in the ewe (Evans, Dahl, Mauger, & Karsch, 1995; Evans, Dahl, Mauger, Padmanabhan, et al., 1995), and a pulsatile mode of LH secretion during the preovulatory LH surge is also evident in other species including humans (Rossmanith et al., 1990). Therefore, we believe that our cellular molecular and electrophysiological findings in combination with our computational modelling provide a foundation for understanding the complex role of Kiss1ARH neurons in controlling fertility in the mammal. Finally, our model provides the first comprehensive biophysical description of the conductances underlying the neuronal activity of KissARH neurons, which can serve as a basis for future computational modelling of the KissARH neuronal network and its interactions with other brain regions involved in the complex regulation of mammalian female reproduction.

Methods and Materials

Animals

All the animal procedures described in this study were performed in accordance with institutional guidelines based on National Institutes of Health standards and approved by the Institutional Animal Care and Use Committee at Oregon Health and Science University (OHSU).

Mice

Earlier version: Kiss1Cre transgenic female mice version 2 (Padilla et al., 2018) were selectively bred at OHSU. They also were crossed with heterozygous Ai32 mice (RRID:IMSR_JAX:024109, C57BL/6 background) which carry ChR2 (H134R)–EYFP gene in their Gt(ROSA)26Sor locus (Madisen et al., 2012). All animals were maintained under controlled temperature and photoperiod (lights on at 0600h and off at 1800h) and given free access to food (Lab Diets 5L0D) and water. Where specified, Kiss1Cre mice received viral injections to express channelrhodopsin 2 (ChR2) in Kiss1ARH neurons, fourteen to twenty-one days prior to each experiment as described (Qiu J. et al., 2016). Some of the females were ovariectomized seven days prior to an experiment. Each animal was injected on day 5 following OVX with 0.25 μg E2 or vehicle, followed on day 6 with 1.50 μg E2 or vehicle and used for experiments on day 7 (Bosch et al., 2013).

Revised version: Kiss1Cre transgenic female mice version 2 (Padilla et al., 2018) were selectively bred at OHSU. They also were crossed with heterozygous Ai32 mice (RRID:IMSR_JAX:024109, C57BL/6 background) which carry ChR2 (H134R)–EYFP gene in their Gt(ROSA)26Sor locus (Madisen et al., 2012). All animals were maintained under controlled temperature and photoperiod (lights on at 0600h and off at 1800h) and given free access to food (Lab Diets 5L0D) and water. Where specified, Kiss1Cre mice received viral injections to express channelrhodopsin 2 (ChR2) in Kiss1ARH neurons, fourteen to twenty-one days prior to each experiment as described (Qiu J. et al., 2016). Some of the females were ovariectomized seven days prior to an experiment.

Added in the revised version: Each animal was injected on day 5 following OVX with 0.25 μg E2 or vehicle, followed on day 6 with 1.50 μg E2 or vehicle and used for experiments on day 7 (Bosch et al., 2013).

AAV delivery to Kiss1Cre mice

Earlier version: Fourteen to twenty-one days prior to each experiment, the Kiss1Cremice (>60 d old) received bilateral ARH injections of a Cre-dependent adeno-associated viral (AAV; serotype 1) vector encoding mCherry (AAV1-Ef1α-DIO-mCherry) or AAV1 vectors designed to encode SaCas9 and single-guide RNAs (sgRNAs) (See the SaCas9 section for specifics on the sgRNA design). Using aseptic techniques, anesthetized female mice (1.5% isoflurane/O2) received a medial skin incision to expose the surface of the skull. The glass pipette with a beveled tip (diameter = 45 μm) was filled with mineral oil, loaded with an aliquot of AAV using a Nanoject II (Drummond Scientific). ARH injection coordinates were anteroposterior (AP): −1.20 mm, mediolateral (ML): ± 0.30 mm, dorsoventral (DV): −5.80 mm (surface of brain z = 0.0 mm); 500 nl of the AAV (2.0×1012 particles/ ml) was injected (100 nl/min) into each position, and the pipette left in place for 10 min post-injection, then slowly retracted from the brain. The skin incision was closed using Vetbond (3M) and each mouse received analgesia (Rimadyl, 4-5 mg/kg, s.c.).

Revised version: Fourteen to twenty-one days prior to each experiment, the Kiss1Cre mice (>60 d old) received bilateral ARH injections of a Cre-dependent adeno-associated viral (AAV; serotype 1) vector encoding mCherry (AAV1-Ef1α-DIO-mCherry) or AAV1 vectors designed to encode SaCas9 and single-guide RNAs (sgRNAs) (See the SaCas9 section for specifics on the sgRNA design). Using aseptic techniques, anesthetized female mice (1.5% isoflurane/O2) received a medial skin incision to expose the surface of the skull. The glass pipette with a beveled tip (diameter = 45 μm) was filled with mineral oil, loaded with an aliquot of AAV using a Nanoject II (Drummond Scientific). ARH injection coordinates were anteroposterior (AP): −1.20 mm, mediolateral (ML): ± 0.30 mm, dorsoventral (DV): −5.80 mm (surface of brain z = 0.0 mm); 500 nl of the AAV (2.0×1012 particles/ ml) was injected (100 nl/min) into each position, and the pipette left in place for 10 min post-injection, then slowly retracted from the brain. The skin incision was closed using Vetbond (3M) and each mouse received analgesia (Rimadyl, 4-5 mg/kg, s.c.).

Generation of AAV1-FLEX-SaCas9-U6-sgTrpc5

The generation of AAV1-FLEX-SaCas9-U6-sgTrpc5 viruses were done at the University of Washington using published methods (Gore et al., 2013; Hunker et al., 2020). The constructs of sgRNAs for Trpc5 (AAV1-FLEX-SaCas9-U6sgTrpc5) were designed to target exon2 and exon7, respectively (Figure 11A and B). To achieve Trpc5 mutagenesis in Kiss1ARH neurons, Kiss1Cre mice were co-injected with AAV1-DIO-mCherry, AAV1-FLEX-SaCas9-U6-sgTrpc5-exon2, and AAV1-FLEX-SaCas9-U6-sgTrpc5-exon7 at a ratio of 10%, 45%, and 45%, respectively. Control animals were co-injected with AAV1-FLEX-SaCas9-U6-sgRosa26 and AAV1-DIO-mCherry at a ratio of 90% and 10%, respectively. AAV1-DIO-mCherry was co-injected with Cas9 vectors to confirm the targeting of injections and visualize the infected Kiss1ARH neurons. Each animal was used for experiments three weeks after the viral injection, with OVX performed one week prior to the experiments.

Visualized whole-cell patch recording

Electrophysiological and optogenetic studies were made in coronal brain slices (250 μm) containing the ARH from AAV1-EF1α-DIO-mCherry injected Kiss1Cre:GFP or Kiss1-Cre:EYFP::AI32 mice, which were vehicle-treated OVX, and E2-treated OVX females 10 weeks and older as previously described (Qiu J. et al., 2016; Qiu J. et al., 2018). Whole-cell patch recordings were performed in voltage-clamp and current-clamp as previously described (Qiu J. et al., 2018) using an Olympus BX51 W1 fixed stage scope out-fitted with epifluorescence and IR-DIC video microscopy. Patch pipettes (A-M Systems; 1.5 μm outer diameter borosilicate glass) were pulled on a Brown/Flaming puller (Sutter Instrument, model P-97) and filled with the following solution (unless otherwise specified): 128 mM potassium gluconate, 10 mM NaCl, 1 mM MgCl2, 11 mM EGTA, 10 mM HEPES, 2 mM ATP, and 0.25 mM GTP adjusted to pH 7.3 with KOH; 295 mOsm. Pipette resistances ranged from 3.5–4 MΩ. In whole-cell configuration, access resistance was less than 30 MΩ; the access resistance was 80% compensated. The input resistance was calculated by measuring the slope of the I-V relationship curve between −70 and −50 mV. Standard whole-cell patch recording procedures and pharmacological testing were performed as previously described (Qiu J. et al., 2003; Qiu J. et al., 2014). Electrophysiological signals were digitized with a Digidata 1322A (Axon Instruments) and the data were analyzed using p-Clamp software (Molecular Devices, Foster City, CA). The liquid junction potential was corrected for all data analysis.

Earlier version: For optogenetic stimulation, a light-induced response was evoked using a light-emitting diode (LED) 470 nm blue light source controlled by a variable 2A driver (ThorLabs, Newton, NJ) with the light path directly delivered through an Olympus 40x water-immersion lens. For high-frequency (20 Hz) stimulation the length of stimulation was 10 seconds (light intensity 0.9 mW and pulse duration, 10 ms) (Qiu J. et al., 2016).

Revised version: For optogenetic stimulation, a light-induced response was evoked using a light-emitting diode (LED) 470 nm blue light source controlled by a variable 2A driver (ThorLabs, Newton, NJ) with the light path directly delivered through an Olympus 40x water-immersion lens. For high-frequency (20 Hz) stimulation the length of stimulation was 10 seconds (Qiu J. et al., 2016).

For studying the firing pattern of Kiss1ARH neurons from OVX and E2-treated OVX mice, the electrodes were filled with an internal solution consisting of the following (in mM): 120 KMeSO4, 5 Na2 phosphocreatine, 10 KCl, 2 MgSO4, 10 HEPES, 4 K2ATP, and 0.25 GTP adjusted to pH 7.3 with KOH; 290 mOsm. The bath solution as described (Ashhad and Feldman, 2020) consisted of (in mM): 124 NaCl, 3.5 KCl, 25 NaHCO3, 0.5 NaH2PO4 and 30 D-glucose,1 MgSO4, 1.5 CaCl2, gassed with 95% O2 /5% CO2 (pH 7.4, 310 mOsm).

For studying the activation/inactivation characteristics of the Ca2+ current, the electrodes were filled with an internal solution as described (Qiu J. et al., 2003) consisting of the following (in mM): 100 Cs+ gluconate, 20 TEA-Cl, 10 NaCl, 1 MgCl2, 10 HEPES, 11 EGTA, 1 ATP, 0.25 GTP, the pH was adjusted to 7.3 with CsOH at 300 mOsm. The bath solution as described (Zhang X. B. and Spergel, 2012) consisted of (in mM) 117.5 NaCl, 25 NaHCO3, 1.25 NaH2PO4, 10 TEA-Cl, 2 CaCl2, 1 MgCl2, 20 sucrose and 5 glucose, gassed with 95% O2 /5% CO2 (pH 7.4, 309 mOsm) and supplemented with 1 μM TTX, 10 μM CNQX, 50 μM AP5 and 100 μM picrotoxin. The effects of estrogen treatment on the peak current, peak current density, and activation/inactivation characteristics of calcium current were measured. Activation curves were fitted by the Boltzmann equation: I/Imax=1/{1+exp[V1/2-Vs)/k]}, where I is the peak current at the step potential Vs, Imax is the peak current amplitude, V1/2 is the step potential yielding half-maximum current, and k is the slope factor. Inactivation curves were fit with the Boltzmann equation: I/Imax = 1 - 1/ {1 + exp [(VH - V1/2)/k]}, where I is the peak current at the step potential VH, Imax is the peak current amplitude, V1/2 is the step potential at which half the current is inactivated, and k is the slope factor.

To record M-currents, pipettes were filled with an internal solution consisting of 10 mM NaCl, 128 mM K-gluconate, 1 mM MgCl, 10 mM HEPES, 1 mM ATP, 1.1 mM EGTA, and 0.25 mM GTP (pH 7.3; 290 mOsm). During voltage-clamp, we employed a standard deactivation protocol (Conde and Roepke, 2019; Roepke et al., 2011) to measure potassium currents. This involved 500-ms voltage steps ranging from –30 to –75 mV in 5-mV increments, following a 300-millisecond prepulse to –20 mV. The amplitude of the M-current relaxation or deactivation was quantified as the difference between the initial (<10 ms) and sustained current (>475 ms) of the current trace.

The bath solution for whole-cell recording of BK, SK and M currents was aCSF supplemented with 1 μM TTX, 10 μM CNQX, 50 μM AP5 and 100 μM picrotoxin.

Electrophysiological solutions/drugs

Earlier version: A standard artificial cerebrospinal fluid (aCSF) was used (Qiu J. et al., 2003; Qiu J. et al., 2010). All drugs were purchased from Tocris Bioscience unless otherwise specified. 1 mM TTX (Alomone Labs), 50 mM DL-2-Amino-5-phosphonopentanoic acid sodium salt (AP5), 10 mM 6-Cyano-7-nitroquinoxaline-2,3-dione disodium (CNQX), 100 mM picrotoxin, 10 mM Nifedipine (Sigma), 2 mM ω-conotoxin GVIA (ConoGVIA; Alomone Labs), 0.5 mM ω-conotoxin MVIIC (ConoMVIIC; Alomone Labs, 50 μM ω-agatoxin IVA (AgaIVA; Alomone Labs), 50 μM SNX-482 (Alomone Labs), 10 mM TTA-P2 (TTAP2; Alomone Labs), 200 mM CdCl2 (Cd2+; Sigma), 40 mM XE991 (Alomone Labs), 100 μM Iberiotoxin (Alomone Labs), 500 μM Apamin (Alomone Labs), 100 mM glutamic acid and 1 mM senktide. Stocks (1000×) were prepared in dimethylsulfoxide DMSO (picrotoxin, TTA-P2 and senktide) or water (TTX, AP5, CNQX, ConoGVIA, ConoMVIIC, AgaIVA, SNX-482, CdCl2 and glutamic acid) and stored at −20°C. Aliquots of the stock solutions were stored as appropriate until needed.

Revised version: A standard artificial cerebrospinal fluid (aCSF) was used (Qiu J. et al., 2003; Qiu J. et al., 2010). All drugs were purchased from Tocris Bioscience unless otherwise specified. 1 mM TTX (Alomone Labs), 50 mM DL-2-Amino-5-phosphonopentanoic acid sodium salt (AP5), 10 mM 6-Cyano-7-nitroquinoxaline-2,3-dione disodium (CNQX), 100 mM picrotoxin, 10 mM Nifedipine (Sigma), 2 mM ω-conotoxin GVIA (ConoGVIA; Alomone Labs), 0.5 mM ω-conotoxin MVIIC (ConoMVIIC; Alomone Labs, 50 μM ω-agatoxin IVA (AgaIVA; Alomone Labs), 50 μM SNX-482 (Alomone Labs), 10 mM TTA-P2 (TTAP2; Alomone Labs), 200 mM CdCl2 (Cd2+; Sigma), 40 mM XE991 (Alomone Labs), 100 μM Iberiotoxin (Alomone Labs), 500 μM Apamin (Alomone Labs), 100 mM glutamic acid and 1 mM senktide. Stocks (1000×) were prepared in dimethylsulfoxide DMSO (picrotoxin, TTAP2 and senktide) or water (TTX, AP5, CNQX, ConoGVIA, ConoMVIIC, AgaIVA, SNX-482, CdCl2 and glutamic acid) and stored at −20°C. Aliquots of the stock solutions were stored as appropriate until needed.

Cell harvesting of dispersed Kiss1Cre neurons and real-time quantitative PCR (qPCR)

Cell harvesting and qPCR was conducted as previously described (Bosch et al., 2013). The ARH was microdissected from basal hypothalamic coronal slices obtained from female Kiss1Cre version 2 mice (Padilla et al., 2018) (n = 4-6 animals/group). The dispersed cells were visualized, patched, and then harvested (10 cells/tube) as described previously (Bosch et al., 2013). Briefly, ARH tissue was incubated in papain (7 mg/ml in oxygenated aCSF) for 50 min at 37° C then washed 4 times in low Ca2+ aCSF and two times in aCSF. For cell dispersion, Pasteur pipettes were flame polished to decreasing tip sizes and gentle trituration used to disperse the neurons onto a glass bottom dish. The plated cells were bathed in oxygenated aCSF using a peristaltic pump to keep the cells viable and clear of debris. Healthy cells with processes and a smooth cell membrane were harvested. Pipettes (World Precision Instruments; 1.5 μm outer diameter borosilicate glass) were pulled on a Brown/Flaming puller (Sutter Instrument, model P-87) to a 10 µm diameter tip. The cells were harvested using the XenoWorks Microinjector System (Sutter Instruments, Navato, CA) which provides negative pressure in the pipette and fine control to draw the cell up into the pipette. Cell pools were harvested and stored at -80°C. All cell pools were DNAse-treated using DNase1. cDNA synthesis was performed as previously described (Bosch et al., 2013).

Primers for the genes that encode for low and high voltage-gated calcium channels, TRPC5, Vglut2, large conductance calcium-activated K+ (BKα) channels, small conductance calcium-activated K+ (SK3) channels and GAPDH were designed using Clone Manager software (Sci Ed Software) to cross at least one intron-exon boundary and optimized as previously described using Power Sybr Green method (Bosch et al., 2013). Real-time qPCR controls included neuronal pools without reverse transcriptase (-RT), hypothalamic RNA with RT (+) and without RT (-), as well as water blanks. Standard curves using ARH cDNA were utilized to determine the real-time PCR efficiency (E = 10(−1/m) – 1) (Biosystems, 2006; Pfaffl, 2001). Only primers resulting in efficiencies of 90-100% were used for analysis. Primer sequences, qPCR parameters and efficiency calculations are provided in Table 1.

mRNA expression analysis

qPCR was performed on a Quantstudio 7 Flex Real-Time PCR System (Applied Biosystems) using Power SYBR Green PCR Master Mix (Applied Biosystems) according to established protocols (Bosch et al., 2013). The comparative ΔΔCT method (Livak and Schmittgen, 2001; Pfaffl, 2001; Schmittgen and Livak, 2008) was used to determine values from duplicate samples of 4 µl for the target genes and 2 µl for the reference gene GAPDH in a 20 μl reaction volume containing 1x Power SYBR Green PCR Master Mix and 0.5 μM forward and reverse primers. Three to six 10-cell pools per animal were analyzed from 4-6 animals per group. The relative linear quantity was determined using the 2-ΔΔCT equation (Livak and Schmittgen, 2001; Pfaffl, 2001; Schmittgen and Livak, 2008). Relative mRNA expression level of target genes in Kiss1Cre neurons was obtained by comparing OVX Oil-treated controls to OVX E2-treated animals. The mean Δ CT for the target genes from the OVX Oil-treated control samples was used as the calibrator. The data were expressed as n-fold change in gene expression normalized to the reference gene GAPDH and relative to the calibrator.

Experimental design and Statistical analysis

For the whole-cell patch recording, pharmacological experiments, only one cell was recorded per slice. Two to three slices were analyzed from each Kiss1Cre mouse, with at least 3-5 mice contributing to each group. For cell harvesting of dispersed Kiss1Cre-YFP neurons and qPCR measurements, 10 cells per pool and 3-6 pools from each animal were used, unless otherwise specified. Statistical comparisons between two groups were performed using an unpaired two-tailed Student’s t-test. Comparisons between more than two groups were performed using the repeated measures, multifactorial ANOVA. If a significant interaction was encountered, we then moved to the one-way ANOVA, followed by the multiple range tests as specified in the appropriate figure legends. All data were analyzed using GraphPad Prism version 6. All data are presented as mean ± standard error of the mean (SEM). Differences were considered statistically significant if the probability of error was less than 5%.

Mathematical model and simulation of neuronal behavior

Earlier version: A mathematical model of the Kiss1ARH neuron was developed and calibrated based on our physiological findings. Here, we employed the Hodgkin-Huxley modelling approach (Hodgkin and Huxley, 1952). Accordingly, the equation describing the membrane potential Vm of the Kiss1ARH neuron is given by where Cm is the membrane capacitance and I is the sum of 12 ionic currents: INaT and INaP are the transient and persistent sodium currents, respectively; IA represents the A current; IM represents the M current; ISK and IBK are the potassium currents through the SK and BK channels respectively; Ih the HCN current; IT the T-type calcium current; ICa represents other calcium currents (L-, N-, P/Q-, and R-type); ITRPCS represents Calcium current throught the TRPC5 channel; IGIRK potassium current through the GIRK channels. Finally, Ileak represents the contribution of leak currents. We use the Hodgkin-Huxley formalism to model current dynamics and their dependence on the membrane potential. A complete specification of all currents along with the complete table of model parameters can be found in the supplementary material. Model simulations were conducted in Matlab using the built-in numerical solver (ode45; based on an explicit Runge-Kutta (4, 5) formula). The Matlab code is available at https://git.exeter.ac.uk/mv286/kiss1-arcuate-neuron-model.

Revised version: A mathematical model of the Kiss1ARH neuron was developed and calibrated based on our physiological findings. Here, we employed the Hodgkin-Huxley modelling approach (Hodgkin and Huxley, 1952). Accordingly, the equation describing the membrane potential Vm of the Kiss1ARH neuron is given by where Cm is the membrane capacitance and I is the sum of 12 ionic currents: INaT and INaP are the transient and persistent sodium currents, respectively; represents the A current; represents the M current; ISK and IBK are the potassium currents through the SK and BK channels respectively; Ih the HCN current; IT the T-type calcium current; ICa represents other calcium currents (L-, N-, P/Q-, and R-type); ITRPCS represents Calcium current throught the TRPC5 channel; IGIRK potassium current through the GIRK channels. Finally, Ileak represents the contribution of leak currents. We use the Hodgkin-Huxley formalism to model current dynamics and their dependence on the membrane potential. A complete specification of all currents along with the complete table of model parameters can be found in the supplementary material. Model simulations were conducted in Matlab using the built-in numerical solver (ode45; based on an explicit Runge-Kutta (4, 5) formula). The Matlab code is available at https://git.exeter.ac.uk/mv286/kiss1-arcuate-neuron-model.